# Assignment 5 - Linear Regression with Gradient Descent  

<img src="./images/gradientDescent_animation.gif" alt="Linear Regression w/Gradient Descent" width="600" align='center'/>

## The submition - important notes
### You must follow the rules described here, regarding assignment submission:  
* <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students
* Test and save your assignment - <b>submit the last tested and saved version</b>
* Work on the <b>original notebook</b> of the assignmen 
* do <b>NOT</b> submit an <b>empty assignment</b> 
* do <b>NOT add extra files</b> 
* submit <b>only</b> the <b>notebook (.ipynb)</b> file and <u>NOT a .py/.txt/.rar/.zip</u> file and so on
* <b>do NOT change</b> the notebook file name 

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

## The assignment relates to 2 datasets:
* a simple dataset with 2 columns: 'attr','target'
  * It will be used for training 1D linear regression model
* An additional dataset, re: house prices
  * It will be used for training a multivariant linear regression model
* You will also need to implement other functions, such as a function for evaluation

## Assignment 5 - Linear Regression with Gradient Descent (total possible points: 10 points)
<u>The assignment consists of 6 parts</u>:<br/>
* a preceding step - import packages
* Part 1 - Student information methods ----> needs stundent implementation
* Part 2 - univariate (1D) linear regression with gradient descent - 4 points
  * part 2a - load and display the '2D' dataset ----> run only
  * part 2b - univariate_predict ----> Student's implementation - total 1 point
  * part 2c - univariate_fit ----> Student's implementation - total 3 points
* Part 3 - The Dataset - House pricing - 5 points
  * Part 3a - load the house pricing dataset ----> run only
  * part 3b - scale and de-scale data using t-score (remove mean, divide std) ----> run only
  * Part 3c - add fictive attribute to x_feature_vector ----> run only
  * Part 3d - multivariate linear regression - predict ----> Student's implementation - total 1 point
  * Part 3e - multivariate linear regression training with gradient descent ----> total 4 points
* Part 4 - Evaluation - RMSE - Root Mean Square Error  ----> Student's implementation - total 1 point 

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
import sys
import os
import pathlib 
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import axes3d, Axes3D
import math
import random
get_ipython().run_line_magic('matplotlib', 'inline')
# --------------------------------------------------------

# --------------------------------------------------------
# Note: do NOT add imports, such as sklearn, which could
# answer the following questions.
# --- add your imports here IF NEEDED:


# --------------------------------------------------------
```

<h3>Part 1 - Student information methods:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>  

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myName</b><br/>
<pre>
The following is expected:
--- the method needs to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myName-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myName
# --------------------------------------------------------
def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # YOUR CODE HERE
    raise NotImplementedError()
```

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the 'YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myId-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myId
# --------------------------------------------------------
def myId():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<u>Part 1 Student information method - test</u><br/>
* Test of 'myName' and 'myId' methods implementation  
* It tests the correctness their implementation

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Test:
## the 'myName' and the 'myId' methods implementation 
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

## Part 2 - Dataset 1 - 2D dataset
* includes 2 columns: 
  * attr   - the attribute to use (a singe attribute x1 for every feature vector)
  * target - the actual value (y)  

#### Part 2a - load  dataset 1 ----> run only
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# The following loads the dataset
print ('Loading the multivariate 2D dataset (dataset-1) ...')
df_1d_attr_target = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_train1.csv')
X_vectors = df_1d_attr_target.drop('target',axis=1)
y_categories = df_1d_attr_target['target']
plt.scatter(X_vectors, y_categories)
plt.show()
df_1d_attr_target.head()
```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# describe dataset
df_1d_attr_target.describe()
```

#### Part 2b  - univariate predict ----> Student's implementation - total 1 point

<u>Part 2b - univariate linear regression predict </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>univariate_predict</b><br/>

<pre>
Change the following function to calculate and return the predicted values, for the given model (with w_0, w_1).
This is done using the univariate linear regression model y_hat=w0+w1*x1.
------------
input parameters:
- x_feature_vectors - a dataframe containing all feature vectors for which we want to predict the y_hat values,
              using the univariat (one feature) linear regression model
    note: x_feature_vectors could be the training feature vectors for interim stages of training, or 
          the test feature vectots for real predictions 
- w_0 - the w_0 parameter of the linear regression model
- w_1 - the w_1 parameter of the linear regression model
    note: w_0, w_1 could be interim value, during training, or final values for testing.
------------
return value:
- y_hat - a series of the predictions for each input instance 
  notes:
        * the y_hat prediction vector could be of an intermediate model during training or 
        for prediction of new examples during evaluation
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: univariate_predict-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: univariate_predict
# --------------------------------------------------------
def univariate_predict(x_feature_vectors, w_0, w_1):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

```python
# --------------------------------------------------------
# VISUALIZATION ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Visualization for the:
## 'univariate_predict' method output 
# --------------------------------------------------------
def plotPointsAndLine(x1Series,y_Series,y_Hat):
    # Plot the values
    plt.scatter(x1Series, y_Series, c = 'b', marker='o')
    plt.xlabel('attribute (x1)')
    plt.ylabel('target (y)')
    plt.plot([min(x1Series), max(x1Series)], [min(y_Hat), max(y_Hat)], color='red') 
    plt.show()    
# --------------------------------------------------------

df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
yHat = None
# --------------------------------------------------------
df_1d_attr_target_train = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_train1.csv')
df_1d_attr_target_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_test1.csv')
X_vectors = df_1d_attr_target_train.drop('target',axis=1)
y_categories = df_1d_attr_target_train['target']
yHat = univariate_predict(df_1d_attr_target_train, 5, 1)

x1TrainSeries = X_vectors.iloc[:,-1]
print('show some model and how it fits the data:')
plotPointsAndLine(x1TrainSeries,y_categories,yHat)

#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
#-------------------------
```

<!-- grade_id: test-univariate_predict, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'univariate_predict' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'univariate_predict' method ...\n") 
# --------------------- 

# --------------------------------------------------------
df_1d_attr_target_train = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_train1.csv')
df_1d_attr_target_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_test1.csv')
X_train = df_1d_attr_target_train.drop('target',axis=1)
y_train = df_1d_attr_target_train['target']
X_test = df_1d_attr_target_test.drop('target',axis=1)
y_test = df_1d_attr_target_test['target']
# --------------------------------------------------------

yHat = univariate_predict(X_train, 2, 1)
assert [int(val) for val in yHat.iloc[0:3]]==[43, 34, 52], 'wrong values from linear regression output'

print ("----> This 'univariate_predict' test passed successfully :-) \n")

#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
#-------------------------
```

## Part 2c - univariate linear regression with gradient descent  
<img src="./images/opt_momentum.png" alt="bagging-classifier" width="400" align='center'/>

#### Part 2c - univariate linear regression with gradient descent ----> Student's implementation - total 3 points
<img src="./images/univariate-gradient-descent.png" alt="univariate gradient-descent" width="600" align='center'/>

###### <u>Part 2c - univariate linear regression fit </u> ----> needs stundent implementation - 3 points<br/>
<u>method name</u>: <b>univariate_fit</b><br/>

<pre>
Change the following method to train a univariate (one feature) linear regression model, 
using the gradient descent algorithm.        
------------
input parameters:
- X_train - a dataframe containing all feature vectors of the train set.
- y_train - a series of containing all class values per train instance.
- w_0_initiail - the initial value of w0 
- w_1_initiail - the initial value of w1
- alpha - the learning rate
- num_iterations - number of iterations to run gadient descent    
------------
return values (comma seperated):
- w_0_trained - the trained value of w_o
- w_1_trained - the trained value of w_1
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: univariate_fit-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 3
# --------------------------------------------------------

# --------------------------------------------------------
# method name: univariate_fit
# --------------------------------------------------------
def univariate_fit(X_train,y_train,w_0_initiail,w_1_initiail,alpha,num_iterations):
    ## use partialDerivative_w0_w1_ForTrainset
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test1-univariate_fit, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'univariate_fit' method implementation 
# --------------------------------------------------------

#---------------------
print ("test #1 - basic test of 'univariate_fit' method for instances ...\n") 
#---------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------------------------------------------

# --------------------------------------------------------
df_1d_attr_target_train = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_train1.csv')
df_1d_attr_target_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_test1.csv')
X_train = df_1d_attr_target_train.drop('target',axis=1)
y_train = df_1d_attr_target_train['target']
X_test = df_1d_attr_target_test.drop('target',axis=1)
y_test = df_1d_attr_target_test['target']
# --------------------------------------------------------
trained_w0, trained_w1 = univariate_fit(X_train,y_train,0,0,0.00001,10)
assert (int(trained_w0*1000),int(trained_w1*100))==(11, 57), 'wrong values for univariate_fit'
print ('trained (10 epoch) wo,w1 = (%f,%f), plot on train points:' %(trained_w0, trained_w1))
x1TrainSeries = X_train.iloc[:,-1]
yHat=univariate_predict(X_train, trained_w0, trained_w1)
plotPointsAndLine(x1TrainSeries,y_train,yHat)

print ("\n----> The #1 'univariate_fit' test passed successfully :-) \n")
#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
trained_w0, trained_w1 = None,None
#-------------------------

```

<!-- grade_id: test2-univariate_fit, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'univariate_fit' method implementation 
# --------------------------------------------------------

#---------------------
print ("test #2 - basic test of 'univariate_fit' method for instances ...\n") 
#---------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------------------------------------------

# --------------------------------------------------------
df_1d_attr_target_train = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_train1.csv')
df_1d_attr_target_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'2d_data_test1.csv')
X_train = df_1d_attr_target_train.drop('target',axis=1)
y_train = df_1d_attr_target_train['target']
X_test = df_1d_attr_target_test.drop('target',axis=1)
y_test = df_1d_attr_target_test['target']
# --------------------------------------------------------
trained_w0, trained_w1 = univariate_fit(X_train,y_train,1,5,0.00001,100)
assert (int(trained_w0*100),int(trained_w1*100))==(93, 148), 'wrong values for univariate_fit'
print ('trained (100 epochs) wo,w1 = (%f,%f), plot on train points:' %(trained_w0, trained_w1))
yHat=univariate_predict(X_train, trained_w0, trained_w1)
x1TrainSeries = X_train.iloc[:,-1]
plotPointsAndLine(x1TrainSeries,y_train,yHat)

print ("\n----> The #2 'univariate_fit' test passed successfully :-) \n")
#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
trained_w0, trained_w1 = None,None
#-------------------------

```

<!-- grade_id: test3-univariate_fit-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'univariate_fit' method implementation 
# --------------------------------------------------------

#-------------------------------------------------------- 
print ("test #3 - test of the 'univariate_fit' method, this test is not visble to the students ...\n") 
#-------------------------
```

## Part 3 - The Dataset - House pricing - 5 points
<img src='./images/For-sale-sign.jpg' alt="The Dataset - House pricing" style="width: 150x"/>

## Similar houses should be similar in price
Expected to share similar attributes:<br/>
* Square footage
* Number of bedrooms

### Part 3a - load the house pricing dataset - run only

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

def plotPointsAndPricingLine(x_dataframe,y_Series,y_Hat):
    fig = plt.figure()
    ax = fig.add_subplot(111, projection='3d')
    ax = Axes3D(fig)
    # Generate the values
    x1_vals = x_dataframe.iloc[:, 0]
    x2_vals = x_dataframe.iloc[:, 1]
    y_vals = y_Series.iloc[:]/1000

    ax.scatter(x1_vals, x2_vals, y_vals, c = 'b', marker='o')
    ax.set_xlabel('size (feet^2)')
    ax.set_ylabel('#rooms')
    ax.set_zlabel('Price($K)')
    
    if y_Hat is not None:
        ax.plot([min(x1_vals), max(x1_vals)], [min(x2_vals), max(x2_vals)],  [min(y_Hat/1000), max(y_Hat/1000)], color='red') 

    plt.show()

# --------------------------------------------------------
# The following loads the house pricing dataset
print ('Loading the multivariate house pricing (2 features) dataset (dataset-2) ...')
df_house_pricing = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
X_vectors = df_house_pricing.drop('price',axis=1)
y_categories = df_house_pricing['price']

print ('displaying data: \n')
plotPointsAndPricingLine(X_vectors, y_categories,None) # y_categories)

df_house_pricing.head()
```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

df_house_pricing.describe()
```

### part 3b - scale and de-scale data using t-score (remove mean, divide std) ----> run only
<img src="./images/var_std_in_sample.png" alt="var in sample" style="width: 250px;"/>  

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# method1 name: t_distibution_mean_std
# --------------------------------------------------------
# The following is performed:
# --- get the mean and std of the series
# ------------
# input parameters:
# - series - the input series to be scaled
# ------------
# return values (comma separated):
# - mean_series - the mean value of the series 
# - std_series - the std value of the series (sample std as shown in the above figure, NOT population std)
# --------------------------------------------------------
def t_distibution_mean_std(series): 
    mean_series = series.mean()
    n_minus_1=series.shape[0]-1
    dist_from_mean = series-mean_series
    std_series  = math.sqrt((dist_from_mean*dist_from_mean).sum()/n_minus_1)
    return [mean_series,std_series]

# --------------------------------------------------------
# method2 name: t_distibution
# --------------------------------------------------------
# The following is performed:
# --- get a dataframe consisting of the mean and std of each column in the dataset.
# --- Notes: - the columns of the output dataframes correspond to the columns of the input dataframe
#            - the 1st and 2nd rows (indexes) consist of the mean and std respectivly.
#            - the std values consists of the sample std as shown in the above figure (and NOT the population std).
# ------------
# input parameters:
# - trainset - an input dataframe of the train set including both the X feature vectors and y categories to be scaled
#              note: that the distribution is calculated on the trainset only and used both for train and test sets
#                    for scaling.
# ------------
# return value:
# - mean_std_dataframe - a dataframe including the mean and std for each column 
# --------------------------------------------------------
def t_distibution(trainset):
    mean_std_dataframe = pd.DataFrame(columns=trainset.columns,index=['mean','std'])
    for col in trainset.columns:
        mean_std_dataframe[col]=t_distibution_mean_std(trainset[col])
    return mean_std_dataframe

# --------------------------------------------------------
# method3 name: t_scale_dataframe
# --------------------------------------------------------
# The following is performed:
# --- scale a feature column, using the t-score scale. The t-distribution information is given as an input parameter.
#     The scaling needs to be done for every column in the input dataframe. The function should return
#     a dataframe with the same columns and indexes, but with the scaled values. 
# ------------
# input parameters:
# - dataset - an input dataframe including both the X feature vectors and y categories to be scaled
# - mean_std_dataframe - a dataframe including the mean and std for each column, to be used in scaling
# ------------
# return value:
# - scaled dataset -  a dataframe containing the same columns and indexes, but with the scaled values.
# ---------------------
def t_scale_dataframe(dataset,mean_std_dataframe):
    scaled_dataset = pd.DataFrame(index=dataset.index,columns=dataset.columns)
    for col in dataset.columns:
        scaled_dataset[col] = (dataset[col]-mean_std_dataframe[col]['mean'])/mean_std_dataframe[col]['std']
    return scaled_dataset

# --------------------------------------------------------
# method4 name: t_descale_attr
# ---------------------------
# The following is performed:
# --- de-scale (the opposite of scaling) the values of the input series, using the t-score scale.
# ------------
# input parameters:
# - scaled_series -  a series containing scaled values.
# - mean_std_series - a column from the mean_std_dataframe for descaling.
# ------------
# return value:
# - de_scaled_series - a seires with the same indexes as the input series, containg the descaled values
# --------------------------------------------------------
def t_descale_series(scaled_series,mean_std_series):
    return scaled_series*mean_std_series['std'] + mean_std_series['mean']
    


# --------------------------------------------------------
print ("show usage of 't_scale_dataframe' output validation ...\n")
df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
# ------
mean_std_dataframe = t_distibution(df_house_pricing_train)
# ------

scaled_dataset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)

de_scaled_price_series = t_descale_series(scaled_dataset['price'],mean_std_dataframe['price'])
print ("train 'price' (head) values          ---> %r" %(df_house_pricing_train['price'].head().tolist()))
print ("train 'price' scaled (head) values   ---> %s" %('['+','.join(['%7.3f...' %(price) for price in scaled_dataset['price'].head().tolist()]) +']'))
print ("train 'price' descaled (head) values ---> %r" %(de_scaled_price_series.head().tolist()))
print ("-------------------------\n")

scaled_dataset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)
de_scaled_size_series = t_descale_series(scaled_dataset['size'],mean_std_dataframe['size'])
print ("train 'size' (head) values          ---> %r" %(df_house_pricing_train['size'].head().tolist()))
print ("train 'size' scaled (head) values   ---> %s" %('['+','.join(['%7.3f...' %(price) for price in scaled_dataset['size'].head().tolist()]) +']'))
print ("train 'size' descaled (head) values ---> %r" %(de_scaled_size_series.head().tolist()))
#--------------------------
df_house_pricing = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
train_set,test_set = None,None
mean_std_dataframe = None
scaled_dataset,de_scaled_size_series = None,None
#-------------------------
```

#### Part 3c - add fictive attribute to x_feature_vector ----> run only

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# The following methods returns a copy of the input 
# dataframe, with an additional fictive feature, with a 
# constant value of 1
# ---------------------
# input parameter:
# - x_feature_vectors - the input dataframe of the feature vectors
# - fictive_attr_name - the name of the fictive feature. This feature will be added as a 
#                       first parameter and will be filled with values of 1.
# ------------------
# returned value:
# x_FVs_w_fictive_attr - the copy of the dataframe, with the added column (of
#                                  the fictive feature)
def insert_fictive_attr_to_df(x_feature_vectors,fictive_attr_name):
    x_FVs_w_fictive_attr = x_feature_vectors.copy()
    x_FVs_w_fictive_attr.insert(loc=0,column=fictive_attr_name,value=1)
    return x_FVs_w_fictive_attr

# --------------------------------------------------------
print ("show usage of 'insert_fictive_attr_to_df' output validation ...\n")
fictive_attr_name = 'fictive_0'

df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
# ------
mean_std_dataframe = t_distibution(df_house_pricing_train)
# ------
print ("show dataframe's first few rows (original feature vectors):")
print (X_train.head())
print ('------------------\n')
print ("show dataframe's first few rows (after adding the fictive attribute):")
x_FVs_w_fictive = insert_fictive_attr_to_df(X_train,fictive_attr_name)
print (x_FVs_w_fictive.head())
# --------------------------------------------------------

```

#### Part 3d - multivariate linear regression  - predict ----> Student's implementation - total 1 point

###### <u>Part 3d - multivariate linear regression  - predict </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>predict</b><br/>

<pre>
Change the following method to calulate and return the multivariate linear regression prediction.
 Note: the 'predict' method should be used both for the predictions on the testset, as well as
        interim steps in the training process.     
------------
input parameters:
- x_feature_vectors - a dataframe containing all feature vectors for which we want to predict the y_hat values,
                      using the multivariate linear regression model's weights.
                 Note: it is expected the the x_feature_vectors contains a fictive column of features, with 
                       the constant value of 1, representing the x_0 fictive feature.
- trained_w - the trained linear regression parameters w vector (or weights from some step in the training) .
                 Note: it is expected the the trained_w series contains a fictive column, representing w_0.
                       In this function you do not have to deal with this fact, but only multiple w series with
                       each x feature vector (or with the whole matrix).
------------
return value:
- y_hat - a series of the predictions for each input feature vector instances

---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:


# ------------------------------------------
```

<!-- grade_id: predict-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1 
# --------------------------------------------------------

# --------------------------------------------------------
# method name: predict
# --------------------------------------------------------
def predict(x_featureVectors,trained_w):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# --------------------------------------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote.
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-predict, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict' method implementation 
# --------------------------------------------------------

#---------------------
print ("test #1 - basic test of 'predict' method for instances ...\n") 
#---------------------
fictive_attr_name = 'fictive_0'

df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
# ------
mean_std_dataframe = t_distibution(df_house_pricing_train)
# ------
scaled_trainset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)
scaled_testset  = t_scale_dataframe(df_house_pricing_test,mean_std_dataframe)
scaled_xTrain = scaled_trainset.drop('price',axis=1)
scaled_yTrain = scaled_trainset['price']
scaled_xTest  = scaled_testset.drop('price',axis=1)
scaled_yTest  = scaled_testset['price']
# ------
fictive_attr_name = 'attr_0'
xTest_w_fctv = insert_fictive_attr_to_df(X_test,fictive_attr_name)
scaled_xTest_w_fctv = insert_fictive_attr_to_df(scaled_xTest,fictive_attr_name)
w_vector = pd.Series(data=[300000, 30, 3000], index=xTest_w_fctv.columns)
yHat = predict(xTest_w_fctv,w_vector)
assert (int(yHat.iloc[0]/1000),int(yHat.iloc[2]/1000))==(347, 438), 'wrong values for predict'

print ("displaying house pricing data")
print ("Note - the output linear regression line (based on the given w vector)\n is not optimal and is only used for testing: \n")
plotPointsAndPricingLine(X_test, y_test,yHat) 

print ("\n----> The 'predict' tests passed successfully :-) \n")

#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
train_set,test_set = None,None
mean_std_dataframe = None
scaled_dataset,de_scaled_size_series = None,None
scaled_xTrain,scaled_yTrain = None,None
scaled_testset,scaled_xTest,scaled_yTest = None,None,None
scaled_trainset,scaled_xTrain,scaled_yTrain = None,None,None
#-------------------------
```

### Part 3e - multivariate linear regression training with gradient descent ----> total 4 points 
<img src='./images/multivariate_gradient_descent.png' alt="multivariate gradient-descent for linear regression" style="width: 200x"/>

###### <u>Part 3e - multivariate linear regression training with gradient descent </u> ----> needs stundent implementation - 4 points<br/>
<u>method name</u>: <b>fit</b><br/>

<pre>
Change the following method to train a multivariate (two or more features) linear regression model, 
using the gradient descent algorithm and returns the trained wight vector w.
  Notes: 
        * at the beginning of the method there is a call to the method 'insert_fictive_attr_to_df'. 
          This returns a copy of the input feature vector dataframe, containing a fictive attribute.
          This values of this fictive attribute are always 1. This is done in order to calculate it as
          one of the features, so you do not need to differentiate between the regular coefficients (of the attributes),
          and the free coefficient w_0.
        * the indexes of the weight vector, which you create, should correspond to the column names, of the 
        'X_train_w_fictive_attr' dataframe returned from the 'insert_fictive_attr_to_df' method.
        * note that the initial value of the weight vectors should be 0 for every attribute (including w_0).
  Note also: you should use the predict method for interim prediction, during training.
------------
input parameters:
- X_train - a dataframe containing all feature vectors of the train set.
- y_train - a series of containing all class values per train instance.
- fictive_attr_name - the name of the feature representing the x_0 fictive feature.
         * This parameter is used in the call to the method 'insert_fictive_attr_to_df', described above.    
- alpha - the learning rate
- num_iterations - number of iterations to run multivariate gadient descent
------------
return value:
- trained_w - a series containing the values of the vector of weights (w), 
              note: the indexes of the returned weight vector series, should correspond to the column names,
                    of the 'X_train_w_fictive_attr' dataframe returned from the 'insert_fictive_attr_to_df' method.
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:



# ------------------------------------------
```

<!-- grade_id: fit-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: fit
# --------------------------------------------------------
def fit(X_train,y_train,fictive_attr_name,alpha,num_iterations):
    # add a column, representing fictive x_0, so you could multiply the
    # w vector with the x feature vectors at once (like taught in the class). The
    # value of this fictive attribute, will always be one (so w_0 behaves like a free
    # parameter). Do not change the following call:
    X_train_w_fictive_attr = insert_fictive_attr_to_df(X_train, fictive_attr_name)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-fit, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------
print ("test #1 - test of the 'fit' method ...\n") 

df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
#---
fictive_attr_name = 'attr_0'
xTest_w_fctv = insert_fictive_attr_to_df(X_test,fictive_attr_name)
#---
mean_std_dataframe = t_distibution(df_house_pricing_train)
#---
scaled_trainset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)
scaled_testset  = t_scale_dataframe(df_house_pricing_test,mean_std_dataframe)
#---
scaled_xTrain = scaled_trainset.drop('price',axis=1)
scaled_yTrain = scaled_trainset['price']
scaled_xTest  = scaled_testset.drop('price',axis=1)
scaled_yTest  = scaled_testset['price']
#---
fictive_attr_name = 'attr_0'
scaled_xTest_w_fctv = insert_fictive_attr_to_df(scaled_xTest,fictive_attr_name)
#---

nWeights = X_train.shape[1] + 1
alpha = 0.01
num_iterations=1
trained_w = fit(scaled_xTrain,scaled_yTrain,fictive_attr_name,alpha,num_iterations)
assert int(trained_w[1]*1000)==15, 'wrong values for fit'
print ('trained (1 epochs) w=%r, plot on test points: \n' %(str(trained_w.tolist())))
yHat = predict(scaled_xTest_w_fctv,trained_w)
de_scaled_yHat = t_descale_series(yHat,mean_std_dataframe['price'])
#y_yhat = pd.DataFrame({'y':yTest,'yHat':de_scaled_yHat})
plotPointsAndPricingLine(X_test, y_test,de_scaled_yHat)

print ("\n----> Test #1 of the 'fit' method passed successfully :-) \n")
#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
trained_w0, trained_w1 = None,None
train_set,test_set = None,None
mean_std_dataframe = None
scaled_dataset,de_scaled_size_series = None,None
scaled_xTrain,scaled_yTrain = None,None
#-------------------------
```

<!-- grade_id: test2-fit, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

#-------------------------------------------------------- 
print ("test #2 - test of the 'fit' method ...\n") 
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------
df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
#---
mean_std_dataframe = t_distibution(df_house_pricing_train)
#---
scaled_trainset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)
scaled_testset  = t_scale_dataframe(df_house_pricing_test,mean_std_dataframe)
#---
scaled_xTrain = scaled_trainset.drop('price',axis=1)
scaled_yTrain = scaled_trainset['price']
scaled_xTest  = scaled_testset.drop('price',axis=1)
scaled_yTest  = scaled_testset['price']
#---
fictive_attr_name = 'attr_0'
xTest_w_fctv = insert_fictive_attr_to_df(X_test,fictive_attr_name)
scaled_xTest_w_fctv = insert_fictive_attr_to_df(scaled_xTest,fictive_attr_name)


alpha = 0.01
num_iterations=10
trained_w = fit(scaled_xTrain,scaled_yTrain,fictive_attr_name,alpha,num_iterations)

assert int(trained_w[2]*1000)==61, 'wrong values for fit'
print ('trained (10 epochs) w=%r\n' %(str(trained_w.tolist())))
yHat = predict(scaled_xTest_w_fctv,trained_w)
de_scaled_yHat = t_descale_series(yHat,mean_std_dataframe['price'])
y_yhat = pd.DataFrame({'y':yTest,'yHat':de_scaled_yHat})
plotPointsAndPricingLine(X_test, y_test,de_scaled_yHat)

num_iterations=100
alpha = 0.01
trained_w = fit(scaled_xTrain,scaled_yTrain,fictive_attr_name,alpha,num_iterations)

print ('trained (100 epochs) w=%r, plot on test points:\n' %(str(trained_w.tolist())))
yHat = predict(scaled_xTest_w_fctv,trained_w)
de_scaled_yHat = t_descale_series(yHat,mean_std_dataframe['price'])
plotPointsAndPricingLine(X_test, y_test,de_scaled_yHat)
print ("\n----> The 'fit' test passed successfully :-) \n")
#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
trained_w0, trained_w1 = None,None
train_set,test_set = None,None
mean_std_dataframe = None
scaled_dataset,de_scaled_size_series = None,None
scaled_xTrain,scaled_yTrain = None,None
#-------------------------*
```

<!-- grade_id: test3-fit-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

#-------------------------------------------------------- 
print ("test #3 - test of the 'fit' method, this test is not visble to the students ...\n") 
```

<!-- grade_id: test4-fit-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

#-------------------------------------------------------- 
print ("test #4 - test of the 'fit' method, this test is not visble to the students ...\n")
```

####  Part 4 - Evaluation - RMSE - Root Mean Square Error  ----> Student's implementation - total 1 point 
<img src="./images/RMSE.png" alt="RMSE" width="300" align='center'/>

###### <u>Part 4 - calculate RMSE (Root Mean Square Error) </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_rmse</b><br/>

<pre>
Change the following method to check the RMSE (root mean squared error) as shown in the above image,
  for running the predicted linear regression values compared to the actual test values.      
------------
input parameters:
- y_test - the actual expected values
- y_hat  - the predicted expected values    
------------
return value:
- mse_val - mse (mean squared error) value
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: calc_rmse-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_rmse
# --------------------------------------------------------
def calc_rmse(y,yPredicted):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-calc_rmse, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_rmse' method implementation 
# --------------------------------------------------------
print ("check basic 'calc_rmse' output validation ...")

df_house_pricing_train = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_train1.csv')
df_house_pricing_test  = pd.read_csv('.'+os.sep+'data'+os.sep+'house_pricing_dataset_test1.csv')
X_train = df_house_pricing_train.drop('price',axis=1)
y_train = df_house_pricing_train['price']
X_test  = df_house_pricing_test.drop('price',axis=1)
y_test  = df_house_pricing_test['price']
#---
mean_std_dataframe = t_distibution(df_house_pricing_train)
#---
scaled_trainset = t_scale_dataframe(df_house_pricing_train,mean_std_dataframe)
scaled_testset  = t_scale_dataframe(df_house_pricing_test,mean_std_dataframe)
#---
scaled_xTrain = scaled_trainset.drop('price',axis=1)
scaled_yTrain = scaled_trainset['price']
scaled_xTest  = scaled_testset.drop('price',axis=1)
scaled_yTest  = scaled_testset['price']
#---
fictive_attr_name = 'attr_0'
xTest_w_fctv = insert_fictive_attr_to_df(X_test,fictive_attr_name)
scaled_xTest_w_fctv = insert_fictive_attr_to_df(scaled_xTest,fictive_attr_name)
#---

nWeights = X_train.shape[1] + 1
alpha = 0.03
num_iterations=10
trained_w = fit(scaled_xTrain,scaled_yTrain,fictive_attr_name,alpha,num_iterations)
yHat = predict(scaled_xTest_w_fctv,trained_w)
de_scaled_yHat = t_descale_series(yHat,mean_std_dataframe['price'])
# calculate rmse in k$:
rmse_descaled = calc_rmse(y_test/1000,de_scaled_yHat/1000)
rmse_scaled   = calc_rmse(scaled_yTest,yHat)

print('trained (10 epochs) rmse_descaled (rmse in k$): %5.1f k$' %(rmse_descaled))
print('trained (10 epochs) rmse_scaled: %5.3f' %(rmse_scaled))
print ('trained (10 epochs) w=%r, plot on test points: \n' %(str(trained_w.tolist())))
plotPointsAndPricingLine(X_test, y_test,de_scaled_yHat)


num_iterations=100
alpha = 0.03
trained_w = fit(scaled_xTrain,scaled_yTrain,fictive_attr_name,alpha,num_iterations)
yHat = predict(scaled_xTest_w_fctv,trained_w)
de_scaled_yHat = t_descale_series(yHat,mean_std_dataframe['price'])
# calculate mse in k$:
rmse_descaled = calc_rmse(y_test/1000,de_scaled_yHat/1000)
rmse_scaled   = calc_rmse(scaled_yTest,yHat)
assert int(rmse_descaled)==67, 'wrong values for calc_rmse'
print ('trained (100 epochs) w=%r, plot on test points:\n' %(str(trained_w.tolist())))
print('trained (100 epochs) rmse_descaled (rmse in k$): %5.1f k$' %(rmse_descaled))
print('trained (100 epochs) rmse_scaled: %5.3f' %(rmse_scaled))

plotPointsAndPricingLine(X_test, y_test,de_scaled_yHat)

# -------
print ("----> The 'calc_rmse' test passed successfully :-) \n")
print ('note: Additional Tests might be executed on our side')  

#--------------------------
df_1d_attr_target = None
X_vectors, y_categories = None,None
xTrain,xTest,yTrain,yTest = None,None,None,None
x1TrainSeries,yHat = None,None
#-------------------------
```