# Assignment 1 - Pandas dataframes Scaling

Motivation - Iris flowers:
<table align='left'>
    <tr><th>Iris Setosa</th><th>Iris Virginica</th><th>Iris Versicolor</th></tr>
    <tr>
        <td><img src="./images/Iris-versicolor.jpg" alt="Iris Versicolor" style="width: 200px;"/></td>
        <td><img src="./images/iris_setosa.jpg" alt="Iris Setosa" style="width: 200px;"/></td>
        <td><img src="./images/iris_virginica.jpg" alt="Iris Virginica" style="width: 200px;"/></td>
    </tr>
</table>    

<h3><u>Feature Set</u></h3>
<table align='left' width="200">
    <tr><td>
        <img src="./images/iris_petal_sepal.png" alt="Iris" width="200" align='left'/>
    </td></tr>
    <tr><td align='left'>
        <div style="text-align:left;font-size:125%"><u>Dataset includes 4 features</u>:</div><br/>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>sepal-length</li>
            <li>sepal-width</li>
            <li>petal-length</li>
            <li>petal-width</li>
        </ul>
    </td></tr>
</table>    

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

## Assignmet 1 - Scaling (total of 10 points)
<u>The assignment consists of 5 parts</u>:<br/>
* preceding step - import packages
* part 1 - Student information methods ----> needs stundent implementation
* Part 2 - load the Iris dataset ----> run only
* Part 3 - column statistics  ----> needs stundent implementation - 4 point
  * Part 3a - mean ----> needs stundent implementation - total 2 points
  * Part 3b - std ---->  needs stundent implementation - total 2 points
* part 4 - scaling  ----> needs stundent implementation - 6 point
  * Part 4a - t-distribution standardization ----> needs stundent implementation - total 3 points
  * Part 4b- minmax normalization ----> needs stundent implementation - total 3 points

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>
<u>Exteranl packages include</u>: 
* pandas - which we use mainly for dataframes and series
* numpy - which we use for advance operations, such as unique values of arrays

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
import os
import pandas as pd
```

### Part 1 - Student information methods:
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b>

# <u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myName</b><br/>
<pre>
The following is expected:
--- the method needs to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myName-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myName
# --------------------------------------------------------
def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # YOUR CODE HERE
    return 'Hen Haim Someh'
```

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the 'YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myId-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myId
# --------------------------------------------------------
def myId():
    # YOUR CODE HERE
    return 313468654
```

<u>Part 1 Student information method - test</u><br/>
* Test of 'myName' and 'myId' methods implementation  
* It tests the correctness their implementation

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Test:
## the 'myName' and the 'myId' methods implementation 
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

Output:
```
'myName' method validation is successfull!
your name is: Hen Haim Someh
--------------------------------------
'myId' method validation is successfull!
your id is: 313468654
```

### Part 2 - Loading dataset - run only
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>



```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# load dataframe from csv:
# --------------------------------------------------------
irisCsvFileName = 'data' + os.sep + 'iris_fearures.csv'
df_iris_features = pd.read_csv(irisCsvFileName)

# --------------------------------------------------------
# display the feature information:
# --------------------------------------------------------
print ('\nThere are %d features in the feature set (each feature vector has a value for every feature)' %(df_iris_features.shape[1]))
featureNames = [col for col in df_iris_features.columns]
print('Feature names:')
print(featureNames)
# --------------------------------------------------------
# display the first few rows of the dataset vectors:
# --------------------------------------------------------
print('\nThe dataset includes %d instances (aka feature vectors)' %(df_iris_features.shape[0]))
print('first few instances:')
df_iris_features.head()
# --------------------------------------------------------
```

Output:
```
There are 4 features in the feature set (each feature vector has a value for every feature)
Feature names:
['sepal-length', 'sepal-width', 'petal-length', 'petal-width']

The dataset includes 150 instances (aka feature vectors)
first few instances:
   sepal-length  sepal-width  petal-length  petal-width
0           5.1          3.5           1.4          0.2
1           4.9          3.0           1.4          0.2
2           4.7          3.2           1.3          0.2
3           4.6          3.1           1.5          0.2
4           5.0          3.6           1.4          0.2
```

## Part 3 - column statistics (mean and std) ----> needs stundent implementation - 4 point

<u>Part 3 - statistics (3a - mean)</u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>calc_column_mean</b><br/>

<pre>
The following is expected:
- write a function which calculate the mean value of a specific column 
------------
input parameters:
- 'dataframe' - the input dataframe, including the column to be scaled
- 'colName' - the column name, of the column which should be scaled (describes column name from input)
------------
return value:
- the mean value of the input column
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_column_mean-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_column_mean
# --------------------------------------------------------
def calc_column_mean(dataframe, colName):
    # YOUR CODE HERE
    return dataframe[colName].sum() / dataframe[colName].size
```

<u>Part 3a mean - <b>test</b></u> (2 points)<br/>
* Test of the 'calc_column_mean' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: test-calc_column_mean, points: 2, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_column_mean' method implementation 
# --------------------------------------------------------

print ("test implementation of the 'calc_column_mean' method ...") 

irisCsvFileName = 'data' + os.sep + 'iris_fearures.csv'
df_iris_features = pd.read_csv(irisCsvFileName)
col0 = df_iris_features.columns[0]

# --------------------------------------------------------
print ("checking for the mean value of the first column ... ")
meanCol0=calc_column_mean(df_iris_features, col0)
assert int (100*meanCol0) == 584, "mean value is not calculated successfully"
print ("\tmean value: %f" %(meanCol0))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_column_mean' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
```

Output:
```
test implementation of the 'calc_column_mean' method ...
checking for the mean value of the first column ... 
	mean value: 5.843333
----> The 'calc_column_mean' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<u>Part 3 - statistics (3b - std)</u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>calc_column_std</b><br/>

<pre>
The following is expected:
- write a function which calculate the standard diviation (std) value of a specific column 
------------
input parameters:
- 'dataframe' - the input dataframe, including the column to be scaled
- 'colName' - the column name, of the column which should be scaled (describes column name from input)
------------
return value:
- the std value of the input column
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_column_std-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_column_std
# --------------------------------------------------------
def calc_column_std(dataframe, colName):
    # YOUR CODE HERE
    std = 0
    mean = calc_column_mean(dataframe, colName)
    
    for x_value in dataframe[colName]:
        std = std + ((x_value - mean)**2)
    
    std = (std / (dataframe[colName].size - 1))**0.5
    return std
    
```

<u>Part 3b std - <b>test</b></u> (2 points)<br/>
* Test of the 'calc_column_std' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: test-calc_column_std, points: 2, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_column_std' method implementation 
# --------------------------------------------------------

print ("test implementation of the 'calc_column_std' method ...") 

irisCsvFileName = 'data' + os.sep + 'iris_fearures.csv'
df_iris_features = pd.read_csv(irisCsvFileName)
col0 = df_iris_features.columns[0]

# --------------------------------------------------------
print ("checking for the std value of the first column ... ")
stdCol0=calc_column_std(df_iris_features, col0)
assert int (100*stdCol0) == 82, "std value is not calculated successfully"
print ("\tstd value: %f" %(stdCol0))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_column_std' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
```

Output:
```
test implementation of the 'calc_column_std' method ...
checking for the std value of the first column ... 
	std value: 0.828066
----> The 'calc_column_std' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

## Part 4 -  scaling ----> needs stundent implementation - 6 point
<img src="./images/architect-scale.jpg" alt="Scale" style="width: 500px;" align='left'/>

### Step 4a - t-distribution Standardizion ----> needs stundent implementation - 3 point
t-distribution Standardization refers to shifting the distribution of each attribute to have a mean of zero and a standard deviation of one (unit variance).

<u>Part 4 - Scaling (4a - t-distribution Standardizion)</u>:<br/>
<u>method name</u>: <b>standardizeColumn</b><br/>

<pre>
The following is expected:
- write a function which scales a specific colum in a dataframe, using the t-distribution Standardizion.
- the function will recieve the dataframe, including the column to be scaled
- the function will return a dataframe including a scaled column
------------
input parameters:
- 'dataframe' - the input dataframe, including the column to be scaled
- 'colName' - the column name, of the column which should be scaled (describes column name from input)
- 'standardizedColName' - the column name, of the scaled column  (describes the new output column name)
------------
return value:
- a dataframe structure containing new column with the scaled data (as well as the original columns)
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: standardizeColumn-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: standardizeColumn
# --------------------------------------------------------
def standardizeColumn(dataframe, colName, standardizedColName):
    colMean = calc_column_mean(dataframe, colName)
    colStd = calc_column_std(dataframe, colName)
    
    dataframe[standardizedColName] = (dataframe[colName] - colMean) / colStd
    return dataframe
```

<u>Part 4a - t-distribution Standardizion scaling - <b>test</b></u> (3 points)<br/>
* Test of the 'standardizeColumn' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: test-standardizeColumn, points: 3, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_column_std' method implementation 
# --------------------------------------------------------

print ("test implementation of the 'standardizeColumn' method ...") 

irisCsvFileName = 'data' + os.sep + 'iris_fearures.csv'
df_iris_features = pd.read_csv(irisCsvFileName)
col0 = df_iris_features.columns[0]
# --------------------------------------------------------
print ("checking if a new column was created ...")
df_iris_features_scaled = standardizeColumn(df_iris_features,col0,col0+'_standardized')
assert col0+'_standardized' in df_iris_features_scaled.columns, 'standardization column was NOT created successfully'
print ("\t standardization column created successfully")
print ("checking the first three standardized values of the first column ... ")
assert df_iris_features_scaled[col0+'_standardized'].apply(lambda x: float(int(x*1000))/1000).head().tolist()[:3]==[-0.897, -1.139, -1.38], "values were NOT standardizatized successfully"
print ("\tvalues were standardtized successfully")
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'standardizeColumn' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
# --------------------------------------------------------
df_iris_features[[col0,col0+'_standardized']].head()

```

Output:
```
test implementation of the 'standardizeColumn' method ...
checking if a new column was created ...
	 standardization column created successfully
checking the first three standardized values of the first column ... 
	values were standardtized successfully
----> The 'standardizeColumn' test passed successfully :-) 

note: Additional Tests might be executed on our side
   sepal-length  sepal-length_standardized
0           5.1                  -0.897674
1           4.9                  -1.139200
2           4.7                  -1.380727
3           4.6                  -1.501490
4           5.0                  -1.018437
```

### Step 4b - Minmax Normalization ----> needs stundent implementation - 3 point
Minmax Normalization refers to rescaling real valued numeric attributes into the range 0 and 1.

<u>Part 4 - Scaling (4b - Minmax Normalization)</u>:<br/>
<u>method name</u>: <b>normalizeColumn</b><br/>

<pre>
The following is expected:
- write a function which scales a specific colum in a dataframe, using the minmax normalization.
- the function will recieve the dataframe, including the column to be scaled
- the function will return a dataframe including a scaled column
------------
input parameters:
- 'dataframe' - the input dataframe, including the column to be scaled
- 'colName' - the column name, of the column which should be scaled (describes column name from input)
- 'normalizedColName' - the column name, of the scaled column  (describes the new output column name)
------------
return value:
- a dataframe structure containing new column with the scaled data (as well as the original columns)
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: normalizeColumn-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 3
# --------------------------------------------------------

# --------------------------------------------------------
# method name: normalizeColumn
# --------------------------------------------------------
def normalizeColumn(dataframe, colName, normalizedColName):
    maxValue = dataframe[colName].max()
    minValue = dataframe[colName].min()
    dataframe[normalizedColName] = (dataframe[colName] - minValue) / (maxValue - minValue)
    return dataframe
```

<u>Part 4a - Minmax Normalization scaling - <b>test</b></u> (3 points)<br/>
* Test of the 'normalizeColumn' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: normalizeColumn-test, points: 3, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

print ("test implementation of the 'normalizeColumn' method ...") 

irisCsvFileName = 'data' + os.sep + 'iris_fearures.csv'
df_iris_features = pd.read_csv(irisCsvFileName)
col2 = df_iris_features.columns[2]

# --------------------------------------------------------
df_iris_features_scaled = normalizeColumn(df_iris_features,col2,col2+'_normalized')
print ("checking if a new column was created ...")
assert col2+'_normalized' in df_iris_features_scaled.columns,'normalization column was NOT created successfully'
print ("\t normalization column created successfully")

print ("checking the first three normalized values of the third column ... ")
assert df_iris_features_scaled[col2+'_normalized'].apply(lambda x: float(int(x*1000))/1000).head().tolist()[:3] == [0.067, 0.067, 0.05], "values were NOT normalized successfully"
print ("\t values normalized successfully")
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'normalizeColumn' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
# --------------------------------------------------------
df_iris_features_scaled[[col2,col2+'_normalized']].head()

```

Output:
```
test implementation of the 'normalizeColumn' method ...
checking if a new column was created ...
	 normalization column created successfully
checking the first three normalized values of the third column ... 
	 values normalized successfully
----> The 'normalizeColumn' test passed successfully :-) 

note: Additional Tests might be executed on our side
   petal-length  petal-length_normalized
0           1.4                 0.067797
1           1.4                 0.067797
2           1.3                 0.050847
3           1.5                 0.084746
4           1.4                 0.067797
```