# Assignment 4 - Bagging (mandatory assignment)
<img src="./images/bags.jpg" alt="The bagging operation" width="400" align='center'/>

## The submition - important notes
### You must follow the rules described here, regarding assignment submission:  
* <b> Work by yourself</b> and submit your own assignment, <b>no pairing</b> to other students
* Test and save your assignment - <b>submit the last tested and saved version</b>
* Work on the <b>original notebook</b> of the assignmen 
* do <b>NOT</b> submit an <b>empty assignment</b> 
* do <b>NOT add extra files</b> 
* submit <b>only</b> the <b>notebook (.ipynb)</b> file and <u>NOT a .py/.txt/.rar/.zip</u> file and so on
* <b>do NOT change</b> the notebook file name 


## The Spam vs. not-Spam Dataset
<img src="./images/spam_or_not_spam.png" alt="Spam vs. not-Spam" width="300" align='center'/>

### About the (Spam vs. not-Spam) dataset:
* The dataset tags email messages as spam or not-spam.
* Classes (=categories): there are two possible classes: not_spam, spam
* Attributes (=features): there are 57 features, including two types of features:
  * word frequencies - the feature name contains a word, with the suffix '_wordFreq' e.g.: 'make_wordFreq'  
    * The word could appear 0 times in a specific email, once or a few times.
    * The values of this feature could be between 0 and 1 (the frequency is relative).
  * Capital Letter pattern attributes - 3 features regarding capital letter patterns. e.g.: 'capitalLet_long' - the length of the longest sequence of capital letters in the email.

### The assignment has to do with the (Spam vs. not-Spam) dataset
* This assignment is a bigger assignment and is graded accordingly (20 points total)
  * It includes different methods to implement, related to the Bagging algorithm 

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

## Assignmet 4 (mandatory) - Bagging (total of 20 points)
<u>The assignment consists of 6 parts</u>:<br/>
* a preceding step - import packages
* Part 1 - Student information methods ----> needs stundent implementation
* Part 2 - load the dataset - total 1 point
  * part 2a - load the 'spam vs not-spam' dataset ----> run only
  * part 2b - change classes values from string to number ----> Student's implementation - total 1 point
* Part 3 - decision stumps classifiction - auxiliary classes and methods - run only
* Part 4 - Training  ----> needs stundent implementation - 10 points
    * Part 4a - bootstrap ----> Student's implementation - total 6 points
    * Part 4b - fit (main bagging train flow) ----> Student's implementation - total 4 points
* Part 5 - Prediction ---> Student's implementation - maximum possible points: 5
* Part 6 - Evaluation - maximum possible points: 4
  * Part 6a - evaluate confusion matrix ----> Student's implementation - total 2 points
  * Part 6b - evaluate precision  ----> Student's implementation - total 1 point
  * Part 6c - evaluate recall     ----> Student's implementation - total 1 point  

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
import sys
import os
import pathlib 
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import random
get_ipython().run_line_magic('matplotlib', 'inline')
```

<h3>Part 1 - Student information methods:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>  

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myName</b><br/>
<pre>
The following is expected:
--- the method needs to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myName-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myName
# --------------------------------------------------------
def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # YOUR CODE HERE
    return "Hen Haim Someh"
```

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the 'YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myId-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myId
# --------------------------------------------------------
def myId():
    # YOUR CODE HERE
    return 313468654
```

<u>Part 1 Student information method - test</u><br/>
* Test of 'myName' and 'myId' methods implementation  
* It tests the correctness their implementation

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: test-myName-myId, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Test:
## the 'myName' and the 'myId' methods implementation 
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

<h3>Part 2 - Loading dataset - total 1 point</h3>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>Part 2a - load the "Spam vs. not Spam" datasets - run only</li>
            <li>Part 2b - change classes values from string to number ----> Student's implementation - total 1 point</li>
        </ul>
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>


#### Part 2a  - load the "Spam vs. not Spam" datasets - run only

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
### The Following section loads the spam vs. not spam dataset
classColName = 'class'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
print ('Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]

print('\t* The train-set includes %d instances (aka feature vectors)' %(X_train.shape[0]))
print('\t* The test-set  includes %d instances (aka feature vectors)' %(X_test.shape[0]))
print ('... done')
print ('--------------------------------------------------------')
# --------------------------------------------------------
# display the spam vs. not spam class information:
# --------------------------------------------------------
spam_classes = np.unique(y_train_classes.values)
print ('There are %d spam classes' %(len(spam_classes)))
print ('spam classes: ' + str(spam_classes))
print ('--------------------------------------------------------')
# --------------------------------------------------------
# display the feature information:
# --------------------------------------------------------
feature_names = [col for col in X_train.columns]
print ('\nThere are %d features in the feature set (each feature vector has a value for every feature)' %(len(feature_names)))
print('Feature names: [' + ','.join(feature_names[:2]) + ',..,' + ','.join(feature_names[-2:]) + ']')
print ('--------------------------------------------------------')
# --------------------------------------------------------
# display the first few rows of the dataset vectors:
# --------------------------------------------------------
print('\nfirst few train feature vectors (and first few features):')
print(X_train.iloc[:,:4].head())
print('-------\n\nfirst few corresponding categories:')
print(y_train_classes.head())
print ('--------------------------------------------------------')
# --------------------------------------------------------

dataset_forTesting = None
X_vectors = None
y_categories = None
# --------------------------------------------------------
```

#### Part 2b  - change classes values from string to number ----> Student's implementation - total 1 point

<u>Part 2b - change classes values from string to number </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>str_categories_to_nums</b><br/>

<pre>
Change the following function to return a pandas series of the categories
with numeric values instead of the original string values.
change the values as following:
-      negative class (string) --> 0
-      positive class (string) --> 1
------------
input parameters:
- y_categories - a series containing the dataset's categories (could be train/test or other categories)
- positive_class_name -   the name of the category (string) of the positive class.
------------
return value:
- y_NumCategories - a series of containing all class values with numeric values per instance.
notes (regarding output):
* each cell value, which doesn't equal the positive_class_name parameter will be considered as negative and 
  recieve a value of 0 in the output series.
* each cell value, which equals the positive_class_name parameter will  be considered as positive and
  recieve a value of 1 in the output series.
* You could NOT assume that the input positive_class_name parameter equals 'spam'    
* The index (of the output series) should be the same as (the index of) the input series.
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: str_categories_to_nums-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: str_categories_to_nums
# --------------------------------------------------------
def str_categories_to_nums(y_categories, positiveClassName):
    # YOUR CODE HERE
    y_NumCategories = []
    
    for i in range(0, len(y_categories)):
        if (y_categories[i] == positiveClassName):
            y_NumCategories.append(1)
        else:
            y_NumCategories.append(0)

    return pd.Series(y_NumCategories, y_categories.index)
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: str_categories_to_nums-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'str_categories_to_nums' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'str_categories_to_nums' method ...\n") 
# --------------------- 

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------- 

classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
print ('Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]

# --------------------- 
y_train_classes_numeric = str_categories_to_nums(y_train_classes, positive_class_name)
# --------------------------------------------------------
assert y_train_classes_numeric is not None, 'y_train_classes_numeric object not initialized'
assert isinstance(y_train_classes_numeric, pd.Series), 'y_train_classes_numeric object is not a series'
assert y_train_classes_numeric.empty == False, 'y_train_classes_numeric dataframe object is empty'
assert sameIndexes(y_train_classes_numeric,y_train_classes), 'y_train_classes_numeric should share the same indexes as y_train_classes'
assert allValidVals(np.unique(y_train_classes_numeric.values),(0,1)), 'classes should be only 0 or 1 (as an integer number)'
num_str_positives     = y_train_classes.where(y_train_classes==positive_class_name).dropna().shape[0]
num_numeric_positives = y_train_classes_numeric.where(y_train_classes_numeric==1).dropna().shape[0]
num_numeric_negatives = y_train_classes_numeric.where(y_train_classes_numeric==0).dropna().shape[0]
num_train_categories =  y_train_classes.shape[0]
assert num_str_positives == num_numeric_positives, 'y_train_classes_numeric should have the same number of positives as y_train_classes' 
assert num_train_categories == num_numeric_positives+num_numeric_negatives, 'method did not work proparly, you must fix this method' 
print ('number of positives: %d, number of negative: %d' %(num_numeric_positives, num_numeric_negatives))

allValidVals = None
sameIndexes = None
classColName = None
positive_class_name = None
folderName = None
datasetNamePrefix = None
train_set = None
test_set = None
X_train = None
y_train_classes = None
X_test = None
y_test_classes = None
num_str_positives = None
num_numeric_positives = None
num_numeric_positives = None
num_numeric_negatives = None
num_train_categories = None
# -------------------------------------------------------- str_categories_to_nums
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'str_categories_to_nums' test passed successfully :-) \n")
```

#### Part 3 decision stumps (and other classification utilities) ----> run only
<img src="./images/treeStumps.jpg" alt="bagging-classifier" width="400" align='center'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------


# --------------------------------------------------------
# auxiliary classes regarding classification
# --------------------------------------------------------


# --------------------------------------------------------
# class 1: DecisionStump
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- wraps an api for decision stumps, so we have a unified api
# ------------
class DecisionStump():
    def __init__(self):
        self._decisionStump = DecisionTreeClassifier(max_depth=1)
        
    def fit(self,X_train,y_train):
        self._decisionStump.fit(X_train,y_train)

    def predict(self,X_test):
        return self._decisionStump.predict(X_test)


# --------------------------------------------------------
# class2: ClassiferInstGen
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- utility classifier to generate objects of the classification algorithm class   
# ------------
class ClassiferInstGen():
    def __init__(self,classifierPyClass):
        self._classifierPyClass = classifierPyClass
        
    def getNewClassifierPyObj(self):
        return self._classifierPyClass()

# --------------------------------------------------------
# class3: BootstrapFeatureClassifer
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- container classifier to save model and bootstrap features    
# ------------
class BootstrapFeatureClassifer():
    def __init__(self,trainModel,featureNames):
        self.model = trainModel
        self.featureNames = featureNames

```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

 
# --------------------------------------------------------
# auxiliary methods regarding classification
# --------------------------------------------------------
        
    
# --------------------------------------------------------
# method name1: trainBootstrapedClassificationModel
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- trains a classification model which matches the sklearn API of fit and predict.
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_sampled_train -  a dataframe containing all feature vectors of the train set
# - y_train - a series of containing all class values per train instance
# Note: the X_sampled_train, y_train parameters could be instance bootsraped, feature bootsraped, both or none
# - classificationAlgo_pyClass - the python 'class' parameter of a classification algorithm 
#                                For instance, the above 'DecisionStump' class.
#                                Note: passing the pyton class is similar to passing a method
#                                      as a parameter. Each call to: classificationAlgo_pyClass()
#                                      creates a new object (also called instance of the class)
#                                      of the type of the class.  
# ------------
# return value:
# - featuresTrainModelObj - an object including trained model and feature names 
# --------------------- 
def trainBootstrapedClassificationModel(X_sampled_train, y_sampled_train,classificationAlgo_pyClass):
    classiferInstGen = ClassiferInstGen(classificationAlgo_pyClass)
    classificationObj = classiferInstGen.getNewClassifierPyObj()
    classificationObj.fit(X_sampled_train, y_sampled_train)
    featuresTrainModelObj = BootstrapFeatureClassifer(classificationObj,X_sampled_train.columns)
    return featuresTrainModelObj

# --------------------------------------------------------
# method name2: trainDecisionStumpModel
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- trains a Decision Stump classification model using the trainBootstrapedClassificationModel.
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_sampled_train -  a dataframe containing all feature vectors of the train set
# - y_sampled_train - a series of containing all class values per train instance
# Note: the X_sampled_train, y_train parameters could be instance bootsraped, feature bootsraped, both or none
# ------------
# return value:
# - featuresTrainModelObj - an object including trained model and feature names 
# --------------------- 
def trainDecisionStumpModel(X_sampled_train, y_sampled_train):
    return trainBootstrapedClassificationModel(X_sampled_train, y_sampled_train,DecisionStump)


# --------------------------------------------------------
# method name3: classifierPredict
# --------------------------------------------------------

# --------------------------------------------------------
# What does the method do?
# --- predict test examples using a classification model (which corresponds to sklearn classifier APIs)
#     - It does so using the sklearn built in method
# ------------
# input parameters:
# - X_test -  a dataframe containing feature vectors of the test set
# - classification_obj - trained classificaion model (which corresponds to sklearn classifier APIs)
# ------------
# return value:
# - yHat - a series of the predictions for each test instance  
# --------------------- trainBootstrapedFeatureModel(X_sampled_train, y_train,classificationAlgo_pyClass)
def classifierPredict(X_test, featuresTrainModelObj):
    X_adapted = X_test[featuresTrainModelObj.featureNames]
    predictions = featuresTrainModelObj.model.predict(X_adapted)
    yHat = pd.Series(data=predictions,index=X_test.index)
    return yHat

# --------------------------------------------------------

```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 

X_training1 = X_train.iloc[:100,:] 
y_training1 = y_train.iloc[:100]
X_training2 = X_train.iloc[100:200,:] 
y_training2 = y_train.iloc[100:200]

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.tree import DecisionTreeClassifier
# --------------------------------------------------------
```

### Usage of auxiliary methods examples
#### You will need to use these 2 methods during the training process:
* trainDecisionStumpModel
* classifierPredict

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# show usages of functions
# --------------------------------------------------------

# --------------------------------------------------------
# example for: the 'trainDecisionStumpModel' method
# You will need this function to train a decision stump model. 
# The following is usage examples:
# --------------------------------------------------------
decisionStumpModel1=trainDecisionStumpModel(X_training1, y_training1)
decisionStumpModel2=trainDecisionStumpModel(X_training2, y_training2)
# --------------------------------------------------------


# --------------------------------------------------------
# example for: the 'classifierPredict' method
# You will need this function to use the trained decision stump model, to predict a new model. 
# The following is usage examples:
# --------------------------------------------------------
print ('A few prediction examples:')
nExamples=6
y_hat1 = classifierPredict(X_test.iloc[:nExamples,:], decisionStumpModel1)
y_hat2 = classifierPredict(X_test.iloc[:nExamples,:], decisionStumpModel2)
for nInd in range(nExamples):
    print ('test instance [%d]: actual: %r, prediction(classifier1): %r, prediction(classifier2): %r' %(nInd,y_test.iloc[nInd],y_hat1.iloc[nInd],y_hat2.iloc[nInd]))

# --------------------------------------------------------
```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat1 = None
y_hat2 = None
X_training1 = None 
y_training1 = None
X_training2 = None 
y_training2 = None
classificationModel1=None
classificationModel2=None

```

### Part 4 - Building a Bagging Classifier
<img src="./images/bagging-classifier.jpeg" alt="bagging-classifier" width="400" align='center'/>

#### Part 4 -  Train a Bagging Classifier
The following cells perform 2 methods:
* part 4a - bootstraping ----> Student's implementation - total 6 points
* part 4b - fit (the bagging main train flow) ----> Student's implementation - total 4 point

#### Part 4a - bootstrap ----> Student's implementation - total 6 points

###### <u>Part 4a - bootstrap </u> ----> needs stundent implementation - 6 points<br/>
<u>method name</u>: <b>bootstrap</b><br/>

<pre>
Change the following method to sample the input dataframe uniformly.
The sampeling could be done on the instances (the rows) or on the features (the columns) and
is dependent on the input boolean parameter do_sample_columns.
Notes:
    * an instanse bootstrap is with replacements 
    * a feature bootstrap is without replacement, but is counted, i.e., if the same feature is selected twice, 
      there will be one less feature in the output. If it selected three times, there will be two less features
      in the output and so on.
    * the ratio regards to the relevant bootstrap type. If it is a feature bootstrap and we have 50 input features,
      and the sample ratio is 0.5, we should expect 25 sampled feature columns the most.
    * if the sample ratio equals 0 you should return None
    * if the sample ratio is greater than one, you need to sample that amount of times. In case of instance selection,
      you will have more feature vectors than the input feature vectors.
        
------------
input parameters:
- X_train - a dataframe containing a training dataset, containing all feature vectors in the train-set.
               * note that the last column (called 'DONT_REMOVE') is a fictive column and it is used only for us.
                 It should always appear in the output, and is not one of the columns being sampled (in 
                 case of feature bootstraping). It should not be counted as one of the columns.
                 In case there are 51 columns in X_train, including 'DONT_REMOVE', and the sample ratio is 10%, 
                 you need to sample 5 (out of 50) columns and include also the 'DONT_REMOVE' column (total 6 columns
                 in the output X_bootstraped in this case of feature bootstraping).
                 It is only used for our test reasons. DO NOT REMOVE IT.
- y_train - a series containing the coresponding (corresponding to X_train) training categories.
- sampleRatio - a floating number, containing the ratio of the sampeling out of feature set,
               * we will derive the number of instances/features from sampleRatio
- do_sample_features -  a boolean variable. 
              * If do_sample_features equals True, you should perform 
                feature bootstrap.
                  * The last column (called 'DONT_REMOVE') is not part of the sampeling 
                    and should not be included.
                     * This means that the last column should always appear, and is not considered
                       a real column,
                       It is only used for trace reasons.
              * If do_sample_features equals False, you should perform instance bootstrap.
------------
return values (comma seperated):
- X_bootstraped - the output bootstraped dataframe.  
              * Note that the last column (called 'DONT_REMOVE') is not part of the sampeling (in 
                 case of feature bootstraping) and should not be counted as one of the columns.
                     * This means that the last column should always appear, and is not considered a real column,
                       It is only used for trace reasons. DO NOT REMOVE IT. (removing it will fail the test)
- y_bootstraped - the corresponding bootstrapted dataframe.
                  * Note that in case of feature bootstraped, the y_train isn't sampled, and could be 
                    returned as y_train (w/o change).

---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: bootstrap-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 6
# --------------------------------------------------------

# --------------------------------------------------------
# method name: bootstrap
# --------------------------------------------------------
def bootstrap(X_train, y_train, sampleRatio, do_sample_features):
    # YOUR CODE HERE
    
    numberOfRows = X_train.shape[0]
    numberofColumns = X_train.shape[1]
   
    if (do_sample_features == False):
        numberOfSamples = int(sampleRatio * numberOfRows)
        X_bootstraped = pd.DataFrame(columns = X_train.columns)
        y_bootstraped = pd.Series()

        for i in range(0, numberOfSamples):
            randomIndex = random.randrange(0 ,numberOfRows - 1)
            X_bootstraped.loc[i] = X_train.iloc[randomIndex]
            y_bootstraped.loc[i] = y_train.iloc[randomIndex]
    
    else: #do_sample_features == True
        columnList = X_train.columns.tolist()
        numberOfSamples = int(sampleRatio * numberofColumns)
        X_bootstraped = pd.DataFrame()
        y_bootstraped = y_train 
        
        for i in range(0, numberOfSamples):
            randomIndex = random.randrange(0 ,numberofColumns - 1)
            X_bootstraped.insert(i, columnList[randomIndex], X_train.iloc[:, randomIndex], allow_duplicates = True)

        X_bootstraped.insert(i + 1, columnList[-1], X_train.iloc[:, -1])     
        X_bootstraped = X_bootstraped.loc[:, ~X_bootstraped.columns.duplicated()]        
 
    return (X_bootstraped, y_bootstraped)
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

```python
# --------------------------------------------------------
# DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
dontRemoveCol = 'DONT_REMOVE'
noteDontRemove = '''Error: the last column ('DONT_REMOVE') was removed (or sampled). 
This column is a fictive column and it is used only for us.
It should always appear in the output, and is not one of the columns being sampled (in 
case of feature bootstraping). It should not be counted as one of the columns.
In case there are 51 columns in X_train, including 'DONT_REMOVE', and the sample ratio is 10%, 
you need to sample 5 (out of 50) columns and include also the 'DONT_REMOVE' column (total 6 columns
in the output X_bootstraped in this case of feature bootstraping).
It is only used for our test reasons. DO NOT REMOVE IT.
'''
# --------------------------------------------------------
```

<!-- grade_id: test-basic-for-instance, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

#---------------------
print ("test #1 - basic test of 'bootstrap' method for instances ...\n") 
#---------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------------------------------------------
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t * Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 
X_train_100 = X_train.iloc[:100,:]
y_train_100 = y_train.iloc[:100]
# ---------------------------------

# --------------------------------------------------------
sampleRatio = 0.5
do_sample_features = False
X_trainSampled, y_trainSampled = bootstrap(X_train_100, y_train_100, sampleRatio,do_sample_features)
# --------------------------------------------------------

# --------------------------------------------------------
assert X_trainSampled is not None, 'X_trainSampled object not initialized'
assert y_trainSampled is not None, 'y_trainSampled object not initialized'
assert isinstance(X_trainSampled, pd.DataFrame), 'X_trainSampled object is not a dataframe'
assert isinstance(y_trainSampled, pd.Series), 'y_trainSampled object is not a series'
assert X_trainSampled.empty == False, 'X_trainSampled dataframe object is empty'
assert y_trainSampled.empty == False, 'y_trainSampled series object is empty'
assert dontRemoveCol in X_trainSampled.columns, noteDontRemove
assert sameIndexes(X_trainSampled,y_trainSampled), 'X_trainSampled should share the same indexes as y_trainSampled'
assert allValidVals(np.unique(y_trainSampled.values),(0,1)), "y_trainSampled's value should be only 0 or 1 (as an integer number)"
assert X_trainSampled.shape[0] == 50, 'X_trainSampled dataframe object should have 50 rows'
assert y_trainSampled.shape[0] == 50, 'y_trainSampled series object should have 50 cells'
assert X_trainSampled.shape[1] == X_train.shape[1], "the features should not be sampled"

print ("----> The basic test of 'bootstrap' for instances passed successfully :-) \n")
# --------------------------------------------------------

# ----------------------
# TEST POSTPROCESS 
# ----------------------
classColName, positive_class_name, folderName, datasetNamePrefix = None, None, None, None
train_set, test_set, y_train_classes, y_test_classes = None, None, None, None
X_train, X_test, y_train, y_test= None, None, None, None
X_train_100, y_train_100 = None, None
sampleRatio, X_trainSampled, y_trainSampled = None, None, None
```

<!-- grade_id: bootstrap-test-basic-for-feature, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

# ---------------------
print ("test #2 - basic test of 'bootstrap' method for features ...\n")
# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------------------------------------------
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t * Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 

# --------------------------------------------------------
columns = X_train.columns
print ("\t * note: we select 51 features, since one (the last one) is fictive, so total 50 to sample from")
columns_50 = [col for col in columns[-51:]]

X_train_50 = X_train.loc[:,columns_50]
sampleRatio = 0.5
do_sample_features = True
X_trainSampled, y_trainSampled = bootstrap(X_train_50, y_train, sampleRatio,do_sample_features)
# --------------------------------------------------------

assert X_trainSampled is not None, 'X_trainSampled object not initialized'
assert y_trainSampled is not None, 'y_trainSampled object not initialized'
assert isinstance(X_trainSampled, pd.DataFrame), 'X_trainSampled object is not a dataframe'
assert isinstance(y_trainSampled, pd.Series), 'y_trainSampled object is not a series'
assert X_trainSampled.empty == False, 'X_trainSampled dataframe object is empty'
assert y_trainSampled.empty == False, 'y_trainSampled series object is empty'
assert dontRemoveCol in X_trainSampled.columns, noteDontRemove
assert X_trainSampled.empty == False, 'X_trainSampled dataframe object is empty'
assert sameIndexes(X_trainSampled,X_train), 'X_trainSampled should share the same indexes as X_train'
assert X_trainSampled.shape[1] >10, 'X_trainSampled dataframe object has a wrong number of sampled features'

print ("----> The basic test of 'bootstrap' for features passed successfully :-) \n")
# --------------------------------------------------------

# ----------------------
# TEST POSTPROCESS 
# ----------------------
classColName, positive_class_name, folderName, datasetNamePrefix = None, None, None, None
train_set, test_set, y_train_classes, y_test_classes = None, None, None, None
X_train, X_test, y_train, y_test= None, None, None, None
X_train_50, columns, columns_50 = None, None, None
sampleRatio, X_trainSampled, y_trainSampled = None, None, None
```

<!-- grade_id: bootstrap-3-compound, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #3 - compound test of the 'bootstrap' method for instances and features ...\n")
# --------------------- 

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
# --------------------------------------------------------
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t* Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 

# --------------------------------------------------------
columns = X_train.columns
print ("\t* we select 51 features, since one (the last one) is fictive, so total 50 to sample from")
columns_50 = [col for col in columns[-51:]]
X_train_100_50 = X_train.loc[:,columns_50].iloc[:100,:]
y_train_100 = y_train.iloc[:100]
# ---------------------------------
sampleRatio = 0.5
do_not_sample_features = False
do_sample_features = True
# ------------
X_train_featureSampled, y_train_featureSampled = bootstrap(X_train_100_50, y_train_100, sampleRatio,do_sample_features)
cols_featureSampled = [col for col in X_train_featureSampled.columns]
# --------------------------------------------------------
assert X_train_featureSampled is not None, 'X_train_featureSampled object not initialized'
assert isinstance(X_train_featureSampled, pd.DataFrame), 'X_train_featureSampled object is not a dataframe'
assert X_train_featureSampled.empty == False, 'X_train_featureSampled dataframe object is empty'
assert y_train_featureSampled.empty == False, 'y_train_featureSampled dataframe object is empty'
assert dontRemoveCol in X_train_featureSampled.columns, noteDontRemove
assert sameIndexes(X_train_featureSampled,X_train_100_50), 'X_train_featureSampled should share the same indexes as X_train_100_50'
assert sameIndexes(X_train_featureSampled,y_train_featureSampled), 'X_train_featureSampled should share the same indexes as y_train_featureSampled'
assert allValidVals(np.unique(y_train_featureSampled.values),(0,1)), "y_train_featureSampled's value should be only 0 or 1 (as an integer number)"
assert X_train_featureSampled.shape[1] >10, 'X_train_featureSampled dataframe object has a wrong number of sampled features'
assert len(cols_featureSampled)==len(list(set(cols_featureSampled))),'duplications in columns after feature sampling'
assert X_train_featureSampled[cols_featureSampled[3]].iloc[2]==X_train_100_50[cols_featureSampled[3]].iloc[2], "'%s' feature value for 3rd row has changed" %(cols_featureSampled[3])
# --------------------------------------------------------

X_train_inst_attr_sampled, y_train_inst_attr_sampled = bootstrap(X_train_featureSampled, y_train_featureSampled, sampleRatio,do_not_sample_features)
cols_inst_attr_sampled = [col for col in X_train_inst_attr_sampled.columns]
# --------------------------------------------------------
assert X_train_inst_attr_sampled is not None, 'X_train_inst_attr_sampled object not initialized'
assert y_train_inst_attr_sampled is not None, 'y_train_inst_attr_sampled object not initialized'
assert isinstance(X_train_inst_attr_sampled, pd.DataFrame), 'X_train_inst_attr_sampled object is not a dataframe'
assert isinstance(y_train_inst_attr_sampled, pd.Series), 'y_train_inst_attr_sampled object is not a series'
assert X_train_inst_attr_sampled.empty == False, 'X_train_inst_attr_sampled dataframe object is empty'
assert y_train_inst_attr_sampled.empty == False, 'y_train_inst_attr_sampled series object is empty'
assert dontRemoveCol in X_train_inst_attr_sampled.columns, noteDontRemove
assert allValidVals(np.unique(y_train_inst_attr_sampled.values),(0,1)), "y_train_inst_attr_sampled's value should be only 0 or 1 (as an integer number)"
assert X_train_inst_attr_sampled.shape[0] == 50, 'X_train_inst_attr_sampled dataframe object should have 50 rows'
assert y_train_inst_attr_sampled.shape[0] == 50, 'y_train_inst_attr_sampled series object should have 50 cells'
assert X_train_inst_attr_sampled.shape[1] == X_train_featureSampled.shape[1], "the features should not be sampled"
for (dont_rem_nRow,nCol) in ((1,5),(5,3),(15,6),(20,9)):
    dont_rem_val=X_train_inst_attr_sampled[dontRemoveCol].iloc[dont_rem_nRow]
    assert X_train_inst_attr_sampled[cols_inst_attr_sampled[nCol]].iloc[dont_rem_nRow]==X_train_100_50.loc[X_train_100_50[dontRemoveCol]==dont_rem_val][cols_inst_attr_sampled[nCol]].iloc[0], "'%s' feature value has changed" %(cols_inst_attr_sampled[nCol])
print ("----> The compound test of 'bootstrap' for instances and features passed successfully :-) \n")
# ----------------------
# TEST POSTPROCESS 
# ----------------------
classColName, positive_class_name, folderName, datasetNamePrefix = None, None, None, None
train_set, test_set, y_train_classes, y_test_classes = None, None, None, None
X_train, X_test, y_train, y_test= None, None, None, None
columns, columns_50, X_train_100_50, y_train_100= None, None, None, None
sampleRatio, do_not_sample_features, do_sample_features= None, None, None
X_train_featureSampled, y_train_featureSampled, X_train_inst_attr_sampled, y_train_inst_attr_sampled = None, None, None, None
```

<!-- grade_id: bootstrap-4-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------
print ("test #4 - test of the 'bootstrap' method, this test is not visble to the students ...\n") 
```

<!-- grade_id: bootstrap-5-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

print ("test #5 - test of the 'bootstrap' method, this test is not visble to the students ...\n") 
```

<!-- grade_id: bootstrap-6-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

print ("test #6 - test of the 'bootstrap' method, this test is not visble to the students ...\n") 
```

#### part 4c - fit (the bagging main flow) ----> Student's implementation - total 4 point:
<img src="./images/bags.jpg" alt="The bagging operation" width="200" align='center'/>

<u>Part 4c -fit (the bagging main flow) </u> ----> needs stundent implementation - 4 points<br/>
<u>method name</u>: <b>fit</b><br/>

<pre>
Change the following function to return the a list of models created 
after bootstraping instances & features.
------------
input parameters:
- X_train - a dataframe containing all feature vectors of the train set
- y_train - a series of containing all class values per train instance
- instanceSampleRatio - the ratio of the sampeling out of training set, 
                        * we will pass it on as a parameter, in order to sample instances.
                        - if instanceSampleRatio<=0, no instance bootstrap is done, and we leave
                          the training instances with no change.
                        - a ratio >= 1 is valid
- featureSampleRatio - the ratio of the sampeling out of feature set,
                        * we will pass it on as a parameter, in order to sample features.
                        - if featureSampleRatio<=0, no feature bootstrap is done, and we leave
                          the features with no change.
                        - a ratio >= 1 is valid

- numModels - number of models to train in bagging
* Note: the X_train, y_train parameters could be instance bootsraped, feature bootsraped or both            
------------
return value:
- a (simple python) list of trained models 
note:
        * the number of models are detrminded by the input parameter 
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: bagging-fit-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 4
# --------------------------------------------------------

# --------------------------------------------------------
# method name: fit
# --------------------------------------------------------
def fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels):
    # note that FOR EACH bootstraped train set and/or feature set you need to create a classifier using:
    # decisionStumpModel=trainDecisionStumpModel(X_train_bootstraped, y_train_corrolated_labels)
    
    # YOUR CODE HERE
    
    trainedModels = []  
    
    for i in range(0, numModels):
        
        if (featureSampleRatio <= 0):
            y_train_corrolated_labels = y_train
            X_train_bootstraped = X_train
        else: # featureSampleRatio > 0
            X_train_bootstraped,y_train_corrolated_labels = bootstrap(X_train, y_train, featureSampleRatio, do_sample_features = True)
            
        if (instanceSampleRatio <= 0):
            y_train_corrolated_labels = y_train
            X_train_bootstraped = X_train
        else: # instanceSampleRatio > 0
            X_train_bootstraped,y_train_corrolated_labels = bootstrap(X_train, y_train, instanceSampleRatio, do_sample_features = False)
        
        decisionStumpModel = trainDecisionStumpModel(X_train_bootstraped, y_train_corrolated_labels)
        trainedModels.append(decisionStumpModel)
    
    return trainedModels
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-fit-1-basic-validity, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #1 - basic validity test of the 'fit' method ...\n")
# --------------------- 

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t* Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 
# --------------------------------------------------------

instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 3
baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
# --------------------------------------------------------

print ("\t * check basic 'fit' output validation ...")
assert baggedModels is not None, 'baggedModels object not initialized'
assert isinstance(baggedModels, list), 'baggedModels object is not a list'
assert None not in baggedModels, 'baggedModels should not include None elements'
assert False not in [isinstance(elem,BootstrapFeatureClassifer) for elem in baggedModels], 'baggedModels should include only DecisionStump objects'

print ("\n----> The #1 test of 'fit' passed successfully :-) \n")
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

<!-- grade_id: bagging-fit-test-2, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #2 - compound test of the 'fit' method ...\n")
# --------------------- 

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------
folderName = '.' + os.sep + 'data' 
datasetCsvFileName = folderName + os.sep + 'spam_notSpam.csv'
# --------------------- 

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t* Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 
# --------------------------------------------------------

instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 3
baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
# --------------------------------------------------------

print ("\t * check basic 'fit' output validation ...")
assert baggedModels is not None, 'baggedModels object not initialized'
assert isinstance(baggedModels, list), 'baggedModels object is not a list'
assert None not in baggedModels, 'baggedModels should not include None elements'
assert False not in [isinstance(elem,BootstrapFeatureClassifer) for elem in baggedModels], 'baggedModels should include only DecisionStump objects'

print ('trying to create inner bagged models prediction ...')
nExamples=20
y_hat_arr = [classifierPredict(X_test.iloc[:nExamples],model) for model in baggedModels]
assert False not in [allValidVals(y_hat.values,(0,1)) for y_hat in y_hat_arr], 'something went wrong with bagging inner models - invalid values found (valid values: 0 or 1)'


print ('\n----------------------')
print ('A few prediction examples:')
for nInd in range(nExamples):
    predictions = '; '.join('[bag-model %d]: %r' %(nPred,y_hat_arr[nPred].iloc[nInd]) for nPred in range(len(y_hat_arr)))
    print ('test instance [%d]: actual: %r, predictions: %r' %(nInd,y_test.iloc[nInd],predictions))

print ("\n----> The #2 test of 'fit' passed successfully :-) \n")

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

<!-- grade_id: bagging-fit-test-3-hidden, points: 2, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'bootstrap' method implementation 
# --------------------------------------------------------

print ("test #3 of the 'fit' method, this test is not visble to the students ...\n") 
```

### Part 5 - Prediction -  maximum possible points: 5
The following cells perform the following:
* step 1 - predict (voting)    ----> Student's implementation - total 3 points

<u>Part 5 -prediction </u> ----> needs stundent implementation - 5 points<br/>
<u>method name</u>: <b>predict</b><br/>

<pre>
Change the following function to return the predicted value for each test instance, 
 based on the bagged trained models. You should return the majority class for
 each test instance.
note: 
    * if there is not majrity, return None for that test instance
      and print a message(this should not happen)
# --------------------------------------------------------
------------
input parameters:
- X_test -  a dataframe containing test feature vectors
- baggingModels - the trained list of bagged models returned from 'fit' 
      (the output of the training process)
------------
return value:
- y_hat - a series of the predictions, with the same index as X_test 
  note:
        * the number of models are detrminded by the input parameter 
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: bagging-predict-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 5
# --------------------------------------------------------

# --------------------------------------------------------
# method name: predict
# --------------------------------------------------------
def predict(X_test, baggingModels):
    # note that FOR EACH model in baggingModels, predict the value using:
    # y_test=classifierPredict(X_test, baggingModel)
    
    # YOUR CODE HERE
    classList = []
    y_hat = pd.Series(X_test.index)
    
    for baggingModel in baggingModels:
        y_test = classifierPredict(X_test, baggingModel)
        classList.append(y_test)
    
    hDataFrame = pd.DataFrame(classList).transpose()
    
    for i in range(0, len(hDataFrame.index)):
  
        count = 0        
        for j in range(0, len(hDataFrame.columns)):
            if (hDataFrame.iloc[i][j] == 1):
                count = count + 1
            elif (hDataFrame.iloc[i][j] == 0):
                count = count - 1 
        
        if (count < 0):
            y_hat[i] = 0  
        elif (count > 0):
            y_hat[i] = 1
        else: # count = 0
            y_hat[i] = 'this should not happen'
      
    return y_hat
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: predict-test-1, points: 3, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #1 - basic validity test of the 'predict' method ...\n")
# --------------------- 
# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t* Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------------------------------------------

print ("\t * check basic 'predict' output validation ...")
instanceSampleRatio = 0.1
featureSampleRatio = 0.5
numModels = 11
baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
yHat = predict(X_test, baggedModels)
assert yHat is not None, 'yHat object not initialized'
assert isinstance(yHat, pd.Series), 'yHat object is not a pandas series'
assert None not in yHat.values, 'yHat should not include None elements'
assert allValidVals(yHat.values,(0,1)), 'prediction err - something went wrong with bagging inner models - invalid values found (valid values: 0 or 1)'

print ('\n----------------------')
print ('A few prediction examples:')
nExamples = 20
for nInd in range(nExamples):
    print ('test instance [%d]: actual: %r, predictions: %r' %(nInd,y_test.iloc[nInd],yHat.iloc[nInd]))

print ("\n----> The #1 test of 'predict' passed successfully :-) \n")

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------

classColName, positive_class_name, folderName, datasetNamePrefix = None, None, None, None
train_set, test_set, y_train_classes, y_test_classes = None, None, None, None
X_train, X_test, y_train, y_test= None, None, None, None
columns, columns_50, X_train_100_50, y_train_100= None, None, None, None
sampleRatio1,sampleRatio2, sampleRatio3, do_not_sample_features, do_sample_features= None, None, None, None, None
X_train_featureSampled1, y_train_featureSampled1, X_train_instSampled1, y_train_instSampled1= None, None, None, None
X_train_featureSampled2, y_train_featureSampled2, X_train_instSampled2, y_train_instSampled2= None, None, None, None
X_train_featureSampled3, y_train_featureSampled3, X_train_instSampled3, y_train_instSampled3= None, None, None, None
X_train_featureSampled, y_train_featureSampled, X_train_inst_attr_sampled, y_train_inst_attr_sampled = None, None, None, None

dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

```

<!-- grade_id: bagging-predict-test-2, points: 2, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #2 - test of the 'predict' method, this test is not visble to the students ...\n") 
# --------------------- 

```

### Part 6 - Evaluation - maximum possible points: 4
The following cells perform the following:
* part 6a - evaluate confusion matrix ----> Student's implementation - total 2 points
* part 6b - evaluate precision  ----> Student's implementation - total 1 point
* part 6c - evaluate recall     ----> Student's implementation - total 1 point

#### Part 6a - evaluate confusion matrix ----> needs stundent implementation - 2 points
<img src="./images/confusion_matrix.jpg" alt="confusion_matrix" width="300" align='center'/>

<u>Part 6a - evaluate confusion matrix </u> ----> needs stundent implementation - 2 points <br/>
<u>method name</u>: <b>calc_confusion_matrix</b><br/>

<pre>
Change the following function to return compute the 4 values of the confusion matrix: TN, FP, FN, TP
Reminder:
* TN - True Negative Count (model predicted 0, actual (test) result 0)
* FP - False Positive Count (model predicted 1, actual (test) result 0)
* FN - False Negative Count (model predicted 0, actual (test) result 1)
* TP - True Positive Count (model predicted 1, actual (test) result 1)
------------
input parameters:
- y_hat -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return values (comma seperated):
- TN - True negatives - number of instances, for which the actual value (from y_test) 
                         is 0 (negative class, not_spam in our dataset) and 
                         the predicted value (from y_hat) is also 0.
- FP - True negatives - number of instances, for which the actual value (from y_test) is 0, but 
                         the predicted value (from y_hat) is 1 (positive class, spam in our dataset).
- FN - False negatives - number of instances, for which the actual value (from y_test) is 1, but                          
                         the predicted value (from y_hat) is 0.
- TP - False negatives - number of instances, for which the actual value (from y_test) is 1 and                          
                         the predicted value (from y_hat) is also 1.
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: confusion-matrix-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_confusion_matrix
# --------------------------------------------------------
def calc_confusion_matrix(y_hat,y_test):
    # YOUR CODE HERE
    TN = len([1 for index in range(0, len(y_test)) if (y_test.iloc[index] == 0) and (y_hat.iloc[index] == 0)])
    FP = len([1 for index in range(0, len(y_test)) if (y_test.iloc[index] == 0) and (y_hat.iloc[index] == 1)])
    FN = len([1 for index in range(0, len(y_test)) if (y_test.iloc[index] == 1) and (y_hat.iloc[index] == 0)])
    TP = len([1 for index in range(0, len(y_test)) if (y_test.iloc[index] == 1) and (y_hat.iloc[index] == 1)])
    
    return TN,FP,FN,TP
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: confusion-matrix-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_confusion_matrix' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #1 - test of the 'calc_confusion_matrix' method, this test is not visble to the students ...\n") 
# --------------------- 

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('\t* Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 

instanceSampleRatio = 0.5
featureSampleRatio = 0.1
numModels = 11
# --------------------------------------------------------

# --------------------------------------------------------
isTN_gt_750=False
isTP_gt_300=False
numTN,numFP,numFN,numTP = 0,0,0,0
nTries = 10
nTry=0
for nTry in range(nTries):
    nTry = nTry + 1
    baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
    y_hat = predict(X_test, baggedModels)
    numTN,numFP,numFN,numTP = calc_confusion_matrix(y_hat,y_test)
    if numTN > 750:
        isTN_gt_750 = True
    if numTP > 300:
        isTP_gt_300 = True
    if isTN_gt_750 and isTP_gt_300:
        break
assert isTN_gt_750 and isTP_gt_300,'the True negatives and True positives is insufficient'
print ('nTry:',nTry)
print ('confusion matrix:')
print ('\tpred-0\tpred-1')
print ('is-0\tTN=%d\tFP=%d' %(numTN,numFP))
print ('is-1\tFN=%d\tTP=%d' %(numFN,numTP))
# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------
print ("\n----> The #1 test of 'calc_confusion_matrix' passed successfully :-) \n")

# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None

yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
```

<!-- grade_id: confusion-matrix-test2-hiden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_confusion_matrix' method implementation 
# --------------------------------------------------------

# --------------------- 
print ("test #2 - test of the 'calc_confusion_matrix' method, this test is not visble to the students ...\n") 
# --------------------- 
```

#### Part 6b - evaluate precision  ----> Student's implementation - total 1 point
<img src="./images/confusion_matrix-precision-recall.jpg" alt="confusion_matrix-precision-recall" width="300" align='center'/>

<u> Part 6b - evaluate precision</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_precision</b><br/>

<pre>
Change the following function to return the precision for class 1 (spam in our dataset)
------------
input parameters:
- y_hat -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value
- precision - float number containing precision value
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: precision-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_precision
# --------------------------------------------------------
def calc_precision(y_hat,y_test):
    # YOUR CODE HERE
    TN,FP,FN,TP = calc_confusion_matrix(y_hat, y_test)
    precision = TP / (TP + FP)
    
    return precision
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: precision-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_precision' method implementation 
# --------------------------------------------------------


# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------- 
print ("test of the 'calc_precision' method, this test is not visble to the students ...\n") 
# --------------------- 
# --------------------------------------------------------
# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 

instanceSampleRatio = 0.5
featureSampleRatio = 0.1
numModels = 5
baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
yHat = predict(X_test, baggedModels)
# --------------------------------------------------------
print ("\t* check basic 'getPrecision' output validation ...")
precision = calc_precision(yHat,y_test)
assert precision is not None, 'precision not initialized'
assert precision>0, 'precision should not be 0'
assert precision<=1, 'precision should not be more than 1 (=100%)'
# --------------------------------------------------------
precision = 0.0
nTries = 10
nTry=0
isPrecision_gt_80=False
for nTry in range(nTries):
    nTry = nTry + 1
    baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
    y_hat = predict(X_test, baggedModels)
    precision = calc_precision(y_hat,y_test)
    if precision>0.8:
        isPrecision_gt_80=True 
        break
assert isPrecision_gt_80,'the precision is too low'

print ('nTry:', nTry)
print ('pricision:', precision)

print ("\n----> The test of 'calc_precision' passed successfully :-) \n")

# --------------------------------------------------------

# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```

#### Part 6c - evaluate recall     ----> Student's implementation - total 1 points
<img src="./images/confusion_matrix-precision-recall.jpg" alt="confusion_matrix-precision-recall" width="300" align='center'/>

<u> Part 6c - evaluate recall</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_recall</b><br/>

<pre>
Change the following function to return the recall for class 1 (spam in our dataset)
------------
input parameters:
- y_hat -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value
- recall - float number containing recall value
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: recall-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_recall
# --------------------------------------------------------
def calc_recall(y_hat,y_test):
    # YOUR CODE HERE
    TN,FP,FN,TP = calc_confusion_matrix(y_hat, y_test)
    recall = TP / (TP + FN)
    
    return recall
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: recall-test, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_recall' method implementation 
# --------------------------------------------------------

# --------------------------------------------------------
allValidVals = lambda arrValues,validVals: False not in (val in validVals for val in arrValues) 
sameIndexes = lambda obj1,obj2: False not in (obj1.index.tolist()[n] == obj2.index.tolist()[n] for n in range(len(obj1.index))) 
innerElementTupple = lambda arrValues: False not in (isinstance(elem,tuple) for elem in arrValues)
# --------------------------------------------------------

# --------------------- 
print ("test #1 - test of the 'calc_recall' method, this test is not visble to the students ...\n") 
# --------------------- 
# --------------------------------------------------------

# --------------------------------------------------------
# Temporary imports:
import sklearn
from sklearn.model_selection import train_test_split
from sklearn.tree import DecisionTreeClassifier
# --------

# --------------------- 
classColName = 'class'
positive_class_name = 'spam'
folderName = '.' + os.sep + 'data' 
datasetNamePrefix = folderName + os.sep + 'spam_notSpam'
# --------------------- 
print ('Loading the spam vs. not spam train and test sets ...')
train_set = pd.read_csv(datasetNamePrefix+'_train1.csv')
test_set  = pd.read_csv(datasetNamePrefix+'_test1.csv')
# --------------------- 
X_train = train_set.drop(classColName,axis=1)
y_train_classes = train_set[classColName]
X_test = test_set.drop(classColName,axis=1)
y_test_classes = test_set[classColName]
# --------------------- 
y_train = str_categories_to_nums(y_train_classes, positive_class_name)
y_test = str_categories_to_nums(y_test_classes, positive_class_name)
# --------------------- 
instanceSampleRatio = 0.5
featureSampleRatio = 0.1
numModels = 11
baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
yHat = predict(X_test, baggedModels)
# --------------------------------------------------------

print ("check basic 'getRecall' output validation ...")
recall = calc_recall(yHat,y_test)
assert recall is not None, 'precision not initialized'
assert recall>0, 'recall should not be 0'
assert recall<=1, 'recall should not be more than 1 (=100%)'
# --------------------------------------------------------
recall = 0.0
nTries = 10
nTry=0
isRecall_gt_50=False
for nTry in range(nTries):
    nTry = nTry + 1
    baggedModels =  fit(X_train, y_train, instanceSampleRatio, featureSampleRatio, numModels)
    y_hat = predict(X_test, baggedModels)
    recall = calc_recall(y_hat,y_test)
    if recall>0.5:
        isRecall_gt_50=True 
        break
assert isRecall_gt_50,'the recall is too low'

print ('nTry:', nTry)
print ('recall:', recall)
# --------------------------------------------------------
print ("\n----> The #1 test of 'calc_recall' passed successfully :-) \n")
# --------------------------------------------------------


# --------------------------------------------------------
if 'datasets' in sys.modules:
    del (datasets)
if 'train_test_split' in sys.modules:
    del (train_test_split)
if 'DecisionTreeClassifier' in sys.modules:
    del (DecisionTreeClassifier)
sys_modules = list(sys.modules.keys())
for mdl in sys_modules:
    if mdl.startswith('sklearn.'):
        del(sys.modules[mdl]) 
del (sklearn)
if 'sklearn' in sys.modules:
    del (sys.modules['sklearn'])
# --------------------------------------------------------

# --------------------------------------------------------
dataset_forTesting = None
X_vectors = None
y_categories = None
X_train = None
X_test = None
y_train = None
y_test = None
y_hat = None
yHat = None

instanceSampleRatio = None
featureSampleRatio = None
numModels = None
classification_PyClass = None
baggedModels = None

nExamples = None
y_hat_arr = None
# --------------------------------------------------------

# --------------------------------------------------------
```