# Assignment 2 - KNN (K-Nearest Neighbors) - our 1st classifier

Motivation - The Titanic dataset: <br/>
<img src="./images/ships-titanic-vehicles-best.jpg" Width="800"/>

<h2><u>Titanic dataset information</u>:</h2><br/>
<table align='left'>
    <tr><td align='left'>
        <div style="text-align:left;font-size:125%"><b><u>Feature Set</u></b>:</div><br/>
        <div style="text-align:left;font-size:115%">Dataset includes <u>7 features</u>:</div><br/>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li><b>Pclass</b> Passenger Class (discrete number: 1 = 1st class; 2 = 2nd class; 3 = 3rd class)</li>
            <li><b>Gender</b> - passanger's age (categorial: 'female', 'male')</li>
            <li><b>Age</b> - passanger's age (continuous number: <= 80)</li>
            <li><b>SibSp</b> - Number of Siblings/Spouses Aboard (discrete number)</li>
            <li><b>Parch</b> - Number of Parents/Children Aboard (discrete number)</li>
            <li><b>fare</b> - Passenger Fare (British pound) - how much did the fare cost (continuous number)</li>
            <li><b>Embarked</b> - Port of Embarkation (categorial: 'C'=Cherbourg; 'Q'=Queenstown; 'S'= Southampton)</li>
        </ul><br/>
        <div style="text-align:left;font-size:125%"><b><u>Categories</u></b>:</div><br/>
        <div style="text-align:left;font-size:115%">Dataset includes <u>2 categories</u>:</div><br/>
        <div style="text-align:left;font-size:115%">- <u>The categories</u>, refer to the survival of the passenger (did they survive?):</div>
        <div style="text-align:left;font-size:115%">- (category values are discrete numbers:) <b>0</b>=No; <b>1</b>=Yes</div><br/><br/>
        <div style="text-align:left;font-size:125%"><b><u>Indexes</u></b>:</div><br/>
        <div style="text-align:left;font-size:115%"><u>The index</u>, refer to the passenger-id:</div><br/>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>(indexes are discrete numbers:) start from <b>1</b></li>
            <li><u>Note</u>: the indexes are <u>NOT</u> row numbers, do not get confused</li>
        </ul>
    </td></tr>
</table>   


## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

## Assignmet 2 - 1st claissifier - KNN (K-Nearest Neighbors) - total of 10 points
<u>The assignment consists of 5 parts</u>:<br/>
* preceding step - import packages
* part 1 - Student information methods ----> needs stundent implementation
* Part 2 - load the Titanic dataset ----> run only
* Part 3 - ditance metrics  ----> needs stundent implementation - 4 point  
  * Part 3a - Manhattan ditance ----> needs stundent implementation 
  * Part 3b - Euclidian ditance ---->  needs stundent implementation
  * Part 3c - Chebyshev ditance ---->  needs stundent implementation
* part 4 - Knn predict  ----> needs stundent implementation - 4 point
  * Part 4a - distance wrapper (vs all train set) ----> run only
  * Part 4b - predict new example ----> needs stundent implementation - total 4 points
  * Part 4a - predict wrapper (predict all testset examples) ----> run only
* part 5 - Evaluation (accuracy) ----> needs stundent implementation - 2 point

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>
<u>Exteranl packages include</u>: 
* pandas - which we use mainly for dataframes and series
* numpy - which we use for advance operations, such as unique values of arrays
* math - for math functions, like sqrt
* os - internal use

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import math 
import os
```

### Part 1 - Student information methods:
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b>

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myName</b><br/>
<pre>
The following is expected:
--- the method needs to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myName-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myName
# --------------------------------------------------------
def myName():
    return "Hen Haim Someh"
```

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the 'YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myId-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myId
# --------------------------------------------------------
def myId():
    return 313468654
```

<u>Part 1 Student information method - test</u><br/>
* Test of 'myName' and 'myId' methods implementation  
* It tests the correctness their implementation

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

<!-- grade_id: test-myName-myId, points: 0, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Test:
## the 'myName' and the 'myId' methods implementation 
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))


## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

Output:
```
'myName' method validation is successfull!
your name is: Hen Haim Someh
--------------------------------------
'myId' method validation is successfull!
your id is: 313468654
```

### Part 2 - Loading dataset  ----> run only
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# load dataframe from csv:
# --------------------------------------------------------

#--------------------------------------------------------
print ('loading Titanic dataset ...')
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
print ('... done')
print ('---------------------------------------')
print ('before split (feature vectors from categories) information:')
print ('number of rows: %d' %(dataframe.shape[0]))
print ('number of columns: %d' %(dataframe.shape[1]))
print ('---------------------------------------')
# --------------------------------------------------------
# display the feature information:
# --------------------------------------------------------
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
print ('after split (feature vectors from categories) information:')
print ('number of features (columns) in the feature set: %d' %(xTitanicInstances.shape[1]))
featureNames = [col for col in xTitanicInstances.columns]
print('Feature names:')
print(featureNames)
print ('\t\t---> each feature vector has a value for every feature')
print ('---------------------------------------')
# --------------------------------------------------------
# display the first few rows of the feature vectors:
# --------------------------------------------------------
print('\nThe dataset includes %d instances (aka feature vectors)' %(xTitanicInstances.shape[0]))
print('----\nfirst few instances:')
print (xTitanicInstances.head())
print ('---------------------------------------')
# --------------------------------------------------------
# display the first few rows of the categories:
# --------------------------------------------------------
print('\nThe dataset includes %d categories (each category corresponds to the parallel feature vector)' %(yTitanicCategories.shape[0]))
print('----\nfirst few categories:')
print (yTitanicCategories.head())
print ('---------------------------------------')
```

Output:
```
loading Titanic dataset ...
... done
---------------------------------------
before split (feature vectors from categories) information:
number of rows: 888
number of columns: 8
---------------------------------------
after split (feature vectors from categories) information:
number of features (columns) in the feature set: 7
Feature names:
['Pclass', 'Gender', 'Age', 'SibSp', 'Parch', 'Fare', 'Embarked']
		---> each feature vector has a value for every feature
---------------------------------------

The dataset includes 888 instances (aka feature vectors)
----
first few instances:
             Pclass  Gender   Age  SibSp  Parch     Fare Embarked
PassengerId                                                      
1                 3    male  22.0      1      0   7.2500        S
2                 1  female  38.0      1      0  71.2833        C
3                 3  female  26.0      0      0   7.9250        S
4                 1  female  35.0      1      0  53.1000        S
5             
```

### Part 3 - (vector) distance metrics ----> needs stundent implementation - 4 point
<img src="./images/distance_metrics.jpg" alt="load dataframe" width="600" align='left'/>

### Part 3 - (vector) distance metrics ----> needs stundent implementation - 4 point
  * Part 3a - Manhattan distance ----> needs stundent implementation - total 1 points (for test)
  * Part 3b - Euclidean distance ----> needs stundent implementation - total 1 points (for test)
  * Part 3b - Chebyshev distance ----> needs stundent implementation - total 1 points (for test)
  * Part 3d - Extra tests - total 1 points

#### Part 3a - Manhattan distance  ----> needs stundent implementation - 1 point
<img src="./images/manhattan.jpg" alt="Manhattan distance" width="350" align='left'/>

<u>Part 3a - Manhattan Distance</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_manhattan_vector_dist</b><br/>

<pre>
The following is expected:
- Change the following function to return the Manhattan distastance between two vectors (series) from the test and train
  ---- the xSeriesTestVector and xSeriesTrainVector are siries and consist of a single row (from train and test correspondigly) 
------------
input parameter:
- xSeriesTestVector  - feature vector (series) from the test set
- xSeriesTrainVector - feature vector (series) from the train set
------------
return value:
- dist - a scalar (number) containing the distance between the two vectors (floating point)
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_manhattan_vector_dist-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_manhattan_vector_dist
# --------------------------------------------------------
def calc_manhattan_vector_dist(xSeriesTestVector,xSeriesTrainVector):
    return np.sum(np.abs(xSeriesTestVector - xSeriesTrainVector))
```

<u>Part 3a - Manhattan Distance - <b>test</b></u> (1 point)<br/>
* Test of the 'calc_manhattan_vector_dist' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell
* additional tests might be performed

<!-- grade_id: test-calc_manhattan_vector_dist, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_manhattan_vector_dist' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_manhattan_vector_dist' method ...") 

# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_manhattan_vector_dist' test passed successfully :-) \n")
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
inst0 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[0]
inst7 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[7]
print (inst0.values)
print (inst7.values)
dist = calc_manhattan_vector_dist(inst0,inst7)
assert dist== 3, 'unexpected manhattan distance: %f' %(dist)
print ('Manhattan Distance calculated successfuly: %d' %(dist))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_manhattan_vector_dist' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
```

Output:
```
testing implementation of the 'calc_manhattan_vector_dist' method ...
----> The 'calc_manhattan_vector_dist' test passed successfully :-) 

[3 1 0]
[3 3 1]
Manhattan Distance calculated successfuly: 3
----> The 'calc_manhattan_vector_dist' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

#### Part 3b - Euclidean distance  ----> needs stundent implementation - 1 point
<img src="./images/euclides.jpg" alt="Euclidean distance" width="300" align='left'/>

<u>Part 3b - Euclidean Distance</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_euclidean_vector_dist</b><br/>

<pre>
The following is expected:
- Change the following function to return the Euclidean distastance between two vectors (series) from the test and train
  ---- the xSeriesTestVector and xSeriesTrainVector are siries and consist of a single row (from train and test correspondigly) 
------------
input parameter:
- xSeriesTestVector  - feature vector (series) from the test set
- xSeriesTrainVector - feature vector (series) from the train set
------------
return value:
- dist - a scalar (number) containing the distance between the two vectors (floating point)
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_euclidean_vector_dist-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_euclidean_vector_dist
# --------------------------------------------------------
def calc_euclidean_vector_dist(xSeriesTestVector,xSeriesTrainVector):
    return np.sqrt(np.sum(np.power((xSeriesTestVector - xSeriesTrainVector),2)))
```

<u>Part 3b - Euclidean Distance - <b>test</b></u> (1 point)<br/>
* Test of the 'calc_euclidean_vector_dist' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell
* additional tests might be performed

<!-- grade_id: test-calc_euclidean_vector_dist, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_euclidean_vector_dist' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_euclidean_vector_dist' method ...") 
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
inst0 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[0]
inst7 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[7]
print (inst0.values)
print (inst7.values)
dist = calc_euclidean_vector_dist(inst0,inst7)
assert float(int(dist*100))/100 == 2.23, 'unexpected euclidean distance: %f' %(dist)
print ('Euclidean Distance calculated successfuly: %f' %(dist))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_euclidean_vector_dist' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
```

Output:
```
testing implementation of the 'calc_euclidean_vector_dist' method ...
[3 1 0]
[3 3 1]
Euclidean Distance calculated successfuly: 2.236068
----> The 'calc_euclidean_vector_dist' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

#### Part 3c - Chebyshev (chessborad) distance  ----> needs stundent implementation - 1 point
<img src="./images/chess_board.jpg" alt="Chebyshev (chessborad) distance" width="200" align='left'/>

<u>Part 3c - Chebyshev Distance</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_chebyshev_vector_dist</b><br/>

<pre>
The following is expected:
- Change the following function to return the Chebyshev distastance between two vectors (series) from the test and train
  ---- the xSeriesTestVector and xSeriesTrainVector are siries and consist of a single row (from train and test correspondigly) 
------------
input parameter:
- xSeriesTestVector  - feature vector (series) from the test set
- xSeriesTrainVector - feature vector (series) from the train set
------------
return value:
- dist - a scalar (number) containing the distance between the two vectors (floating point)
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_chebyshev_vector_dist-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_chebyshev_vector_dist
# --------------------------------------------------------
def calc_chebyshev_vector_dist(xSeriesTestVector,xSeriesTrainVector):
    numberOfItems = xSeriesTestVector.size
    distance = 0
    for i in range (0, numberOfItems):
        distance = np.maximum(abs(xSeriesTestVector[i] - xSeriesTrainVector[i]), distance)
    return distance

```

<u>Part 3c - Chebyshev Distance - <b>test</b></u> (1 point)<br/>
* Test of the 'calc_chebyshev_vector_dist' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell
* additional tests might be performed

<!-- grade_id: test-calc_chebyshev_vector_dist, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_chebyshev_vector_dist' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_chebyshev_vector_dist' method ...") 
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
inst0 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[0]
inst7 = xTitanicInstances[['Pclass','SibSp','Parch']].iloc[7]
print (inst0.values)
print (inst7.values)
dist = calc_chebyshev_vector_dist(inst0,inst7)
assert float(int(dist*100))/100 == 2, 'unexpected chebyshev distance: %f' %(dist)
print ('Chebyshev Distance calculated successfuly: %f' %(dist))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("----> The 'calc_chebyshev_vector_dist' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  
```

Output:
```
testing implementation of the 'calc_chebyshev_vector_dist' method ...
[3 1 0]
[3 3 1]
Chebyshev Distance calculated successfuly: 2.000000
----> The 'calc_chebyshev_vector_dist' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<!-- grade_id: test-hidden-distsance-metrics, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## diffenrent distance methods implementation - extra test
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the distance methods, these tests are not visble to the students ...") 
```

Output:
```
testing implementation of the distance methods, these tests are not visble to the students ...
```

#### Part 4 - KNN - predict  ----> needs stundent implementation - 4 point
<img src="./images/nearest_neighbors.jpg" alt="k-Nearest-neighbors predict" width="300" align='left'/>

### Part 4 - predict (KNN main flow) ----> needs stundent implementation - 4 point
  * Part 4a - calculate distances to training examples (auxiliary function, used by 'predict_one_instance') ----> run only
  * Part 4b - prdict one example ----> needs stundent implementation - total 4 points (for test)
  * Part 4c - predict all test examples (auxiliary function, wraps 'predict_one_instance') ----> run only

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

"""
-----
Part 4a - calculate distances to training examples
--------------------------------------------------------
Auxiliary function: calc_distances
--------------------------------------------------------
- the method wraps the different distance metrics methods
   and calculates the distance between the (single) 
   test feature-vector (series) to all training feature-vectors (dataframe).
   * It does so using the distanceMetric as a parameter 
   (see explaination ahead and usage in the test field).
--------------------------------------------------------
input parameters:
- xSeriesTestVector - is a Series object, from the test set
- xTrainInstances - is a dataframe including all train instances
- distanceMetric - the name of the distance function, not as a string
---- distanceMetric - options are: calc_manhattan_vector_dist ; calc_euclidean_vector_dist ; calc_chebyshev_vector_dist
------------------------
The retuned value is a numpy array, containing distances between the test instance
                and all train instances (ordered by the train instance order) 
"""
# --------------------------------------------------------

#--------------------------------------------------------

# --------------------------------------------------------
# method name: calc_distances
# --------------------------------------------------------
def calc_distances(xSeriesTestVector, xTrainInstances, distanceMetric):
    distances = np.zeros(xTrainInstances.shape[0])
    for i in range(xTrainInstances.shape[0]):
        distances[i] = distanceMetric(xSeriesTestVector, xTrainInstances.iloc[i])
    return distances
```

#### 4b KNN - predict one example  ----> needs stundent implementation - 4 point
<img src="./images/3nn_majority.png" alt="KNN - predict" width="300" align='left'/>

<u>Part 4b KNN - predict one example </u> ----> needs stundent implementation - 3 point<br/>
<u>method name</u>: <b>predict_one_instance</b><br/>

<pre>
The following is expected:
- Change the following function to return the Chebyshev distastance between two vectors (series) from the test and train
  ---- the xSeriesTestVector and xSeriesTrainVector are siries and consist of a single row (from train and test correspondigly) 
------------
input parameter:
- xSeriesTestVector - is a Series object, from the test set to be classified
- xTrainInstances - a dataframe including all train instances to be compared to
- yTrainCategories - a dataframe including all train categories 
- distanceMetric - the name of the distance function, not as a string
   ---- distanceMetric - options are: calc_manhattan_vector_dist ; calc_euclidean_vector_dist ; calc_chebyshev_vector_dist
- k - the number of Nearest Neighbors (we select the majority out of k votes)
------------
return value:
- chosenCategory - a scalar (number) containing the majority category
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: predict-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 3
# --------------------------------------------------------

# --------------------------------------------------------
# method name: predict_one_instance
# --------------------------------------------------------
def predict_one_instance(xSeriesTestVector, xTrainInstances,yTrainCategories,distanceMetric,k):
    distances = calc_distances(xSeriesTestVector, xTrainInstances,distanceMetric)
    # calculate the k- nearest neighbors, based on the above distance array
    # YOUR CODE HERE
    # Zip two lists (distances and yTrainCategories) by distance and sort the results.
    zipArray = zip(distances, yTrainCategories)
    sortedCategoriesList = [i for j, i in sorted(zipArray)] 
    
    categoriesCountArray = np.zeros(2)
    for i in range (0, k):
        categoriesCountArray[sortedCategoriesList[i]] = categoriesCountArray[sortedCategoriesList[i]] + 1
   
    # Check which category ('0' or '1') is more dominant  
    if (categoriesCountArray[0] > categoriesCountArray[1]):
        return 0
    
    return 1

```

<u>Part 4b - Predict one example - <b>test</b></u> (4 points)<br/>
* Test of the 'predict_one_instance' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell
* additional tests might be performed

<!-- grade_id: test-predict_one_instance-manhattan, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict_one_instance' method implementation with a
# Manhattan distance metric
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict_one_instance' method ...") 
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']

xTitanicInstances_5Features = xTitanicInstances[['Pclass','SibSp','Parch','Fare','Age']]
xTitanicInstances_5Features_40 = xTitanicInstances_5Features.iloc[:40]
yTitanicCategories_40 = yTitanicCategories.iloc[:40]
xTitanicInstances_5Features_TestInst = xTitanicInstances_5Features.iloc[-7:-6]

# measure distance without scaling:
manhattanCategory_k3 = predict_one_instance(xTitanicInstances_5Features_TestInst.iloc[-1],xTitanicInstances_5Features_40,yTitanicCategories_40,calc_manhattan_vector_dist,3)
print ('manhattanDist,k=3, inst=-%d, chosen cat: %d' %(7,manhattanCategory_k3))
assert manhattanCategory_k3 == 0,"wrong value for knn, with k=3, dist='manhattan'"


# Feature Scaling
## We will be using standardscaler to transform
from sklearn.preprocessing import StandardScaler
sc = StandardScaler()
## transforming "train_x"
xTrainScaled = pd.DataFrame(sc.fit_transform(xTitanicInstances_5Features_40),index=xTitanicInstances_5Features_40.index,columns=xTitanicInstances_5Features_40.columns)
## transforming "test_x"
xTestScaled = pd.Series(sc.transform(xTitanicInstances_5Features_TestInst)[0],index=xTitanicInstances_5Features_40.columns)
# measure distance with scaling:
manhattanCategory_scaled_k3 = predict_one_instance(xTestScaled,xTrainScaled,yTitanicCategories_40,calc_manhattan_vector_dist,3)
print ('(scaled) manhattanDist,k=3, inst=-%d, chosen cat: %d' %(7,manhattanCategory_scaled_k3))
assert manhattanCategory_scaled_k3 == 1,"wrong value for (scaled) knn, with k=3, dist='manhattan'"
print ('---------------------------------------------------')

# --------------------------------------------------------
# if no exception was thrown by this point, this test was successful:
print ("----> The 'predict_one_instance' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  

```

Output:
```
testing implementation of the 'predict_one_instance' method ...
manhattanDist,k=3, inst=-7, chosen cat: 0
(scaled) manhattanDist,k=3, inst=-7, chosen cat: 1
---------------------------------------------------
----> The 'predict_one_instance' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<!-- grade_id: test-predict_one_instance-euclidian, points: 1, grade: true, solution: false -->
```python
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict_one_instance' method implementation with an
# euclidian distance metric
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict_one_instance' method ...") 
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
xTitanicInstances_5Features = xTitanicInstances[['Pclass','SibSp','Parch','Fare','Age']]
xTitanicInstances_5Features_40 = xTitanicInstances_5Features.iloc[:40]
yTitanicCategories_40 = yTitanicCategories.iloc[:40]
xTitanicInstances_5Features_TestInst = xTitanicInstances_5Features.iloc[-7:-6]

# measure distance without scaling:
euclideanCategory_k3 = predict_one_instance(xTitanicInstances_5Features_TestInst.iloc[-1],xTitanicInstances_5Features_40,yTitanicCategories_40,calc_euclidean_vector_dist,3)
print ('euclideanDist,k=3, inst=-%d, chosen cat: %d' %(7,euclideanCategory_k3))
assert euclideanCategory_k3 == 1,"wrong value for knn, with k=3, dist='euclidean'"

# Feature Scaling
## We will be using standardscaler to transform
from sklearn.preprocessing import StandardScaler
sc = StandardScaler()
## transforming "train_x"
xTrainScaled = pd.DataFrame(sc.fit_transform(xTitanicInstances_5Features_40),index=xTitanicInstances_5Features_40.index,columns=xTitanicInstances_5Features_40.columns)
## transforming "test_x"
xTestScaled = pd.Series(sc.transform(xTitanicInstances_5Features_TestInst)[0],index=xTitanicInstances_5Features_40.columns)
# measure distance with scaling:
euclideanCategory_scaled_k3 = predict_one_instance(xTestScaled,xTrainScaled,yTitanicCategories_40,calc_euclidean_vector_dist,3)
print ('(scaled) euclideanDist,k=3, inst=-%d, chosen cat: %d' %(7,euclideanCategory_scaled_k3))
assert euclideanCategory_scaled_k3 == 1,"wrong value for (scaled) knn, with k=3, dist='euclidean'"
print ('---------------------------------------------------')

# --------------------------------------------------------
# if no exception was thrown by this point, this test was successful:
print ("----> The 'predict_one_instance' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  

```

Output:
```
testing implementation of the 'predict_one_instance' method ...
euclideanDist,k=3, inst=-7, chosen cat: 1
(scaled) euclideanDist,k=3, inst=-7, chosen cat: 1
---------------------------------------------------
----> The 'predict_one_instance' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<!-- grade_id: test-predict_one_instance-chebyshev, points: 1, grade: true, solution: false -->
```python
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict_one_instance' method implementation with an
# chebyshev distance metric
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict_one_instance' method ...") 
fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
xTitanicInstances_5Features = xTitanicInstances[['Pclass','SibSp','Parch','Fare','Age']]
xTitanicInstances_5Features_40 = xTitanicInstances_5Features.iloc[:40]
yTitanicCategories_40 = yTitanicCategories.iloc[:40]
xTitanicInstances_5Features_TestInst = xTitanicInstances_5Features.iloc[-7:-6]

# measure distance without scaling:
chebyshevCategory_k5 = predict_one_instance(xTitanicInstances_5Features_TestInst.iloc[-1],xTitanicInstances_5Features_40,yTitanicCategories_40,calc_chebyshev_vector_dist,5)
print ('chebyshevDist,k=5, inst=-%d, chosen cat: %d' %(7,chebyshevCategory_k5))
assert chebyshevCategory_k5 == 0,"wrong value for knn, with k=5, dist='chebyshev'"

# Feature Scaling
## We will be using standardscaler to transform
from sklearn.preprocessing import StandardScaler
sc = StandardScaler()
## transforming "train_x"
xTrainScaled = pd.DataFrame(sc.fit_transform(xTitanicInstances_5Features_40),index=xTitanicInstances_5Features_40.index,columns=xTitanicInstances_5Features_40.columns)
## transforming "test_x"
xTestScaled = pd.Series(sc.transform(xTitanicInstances_5Features_TestInst)[0],index=xTitanicInstances_5Features_40.columns)
# measure distance with scaling:
chebyshevCategory_scaled_k5 = predict_one_instance(xTestScaled,xTrainScaled,yTitanicCategories_40,calc_euclidean_vector_dist,5)
print ('(scaled) chebyshevDist,k=3, inst=-%d, chosen cat: %d' %(7,chebyshevCategory_scaled_k5))
assert chebyshevCategory_scaled_k5 == 0,"wrong value for (scaled) knn, with k=5, dist='chebyshev'"
print ('---------------------------------------------------')

# --------------------------------------------------------
# if no exception was thrown by this point, this test was successful:
print ("----> The 'predict_one_instance' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  

```

Output:
```
testing implementation of the 'predict_one_instance' method ...
chebyshevDist,k=5, inst=-7, chosen cat: 0
(scaled) chebyshevDist,k=3, inst=-7, chosen cat: 0
---------------------------------------------------
----> The 'predict_one_instance' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<!-- grade_id: test-predict-hidden, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## addtional tests - predict method implementation - extra test
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict' method, these tests are not visble to the students ...") 
```

Output:
```
testing implementation of the 'predict' method, these tests are not visble to the students ...
```

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------


"""
Part 4c - predict all test-set examples using knn.
--------------------------------------------------------
Auxiliary function: predict
--------------------------------------------------------
- the method wraps the 'predict_one_instance' method you
#   implemented, to run the knn prediction on all test examples.
--------------------------------------------------------
input parameters:
- xTestInstances - is adataframe including all test instances
- xTrainInstances - is a dataframe including all train instances
- yTrainCategories - is a series including all train categories
- distanceMetric - the name of the distance function, not as a string
---- distanceMetric - options are: calc_manhattan_vector_dist ; calc_euclidean_vector_dist ; calc_chebyshev_vector_dist
k - number of nearest neighbors
------------------------
The retuned value is a series, containing knn predictions for the corresponding test instances
--------------------------------------------------------
"""
def predict(xTestInstances, xTrainInstances,yTrainCategories,distanceMetric,k):
    yPredicted = pd.Series(index=xTestInstances.index,dtype='float64')
    for i in range(xTestInstances.shape[0]):
        xTestInstance = xTestInstances.iloc[i]
        yPredictedCategory_i = predict_one_instance(xTestInstance,xTrainInstances,yTrainCategories,distanceMetric,k)
        yPredicted.iloc[i]=yPredictedCategory_i
    return yPredicted
```

### Part 5 - evaluation (accuracy)
<img src="./images/accuracy.jpg" alt="Accuracy" width="300" align='left'/>

<u>Part 5 - Evaluation </u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>calc_accuracy</b><br/>

<pre>
The following is expected:
- the method needs to return the accuracy (correctly predicted ratio)
------------
input parameter:
- y_predicted -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value:
------------
- accuracyVal - the accuracy value
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: calc_accuracy-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_accuracy
# --------------------------------------------------------
def calc_accuracy(y_predicted,y_test):
    # YOUR CODE HERE
    accuracyRate = 0
    
    for i in range (0, y_predicted.size):
        if (y_predicted.iloc[i] == y_test.iloc[i]):
            accuracyRate = accuracyRate + 1
            
    accuracyRate = accuracyRate / y_predicted.size
    return accuracyRate
```

<u>Part 5 - Accuracy - <b>test</b></u> (1 point)<br/>
* Test of the 'calc_accuracy' method implementation  

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell
* additional tests might be performed

<!-- grade_id: test-calc_accuracy, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_accuracy' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_accuracy' method ...") 

fileName = 'data' + os.sep + 'titanic_dataset.csv'
dataframe = pd.read_csv(fileName,index_col='PassengerId')
xTitanicInstances = dataframe.drop('Survived',axis=1)
yTitanicCategories = dataframe['Survived']
xTitanicInstances_5Features = xTitanicInstances[['Pclass','SibSp','Parch','Fare','Age']]
xTitanicInstances_5Features_100 = xTitanicInstances_5Features.iloc[:100]
yTitanicCategories_100 = yTitanicCategories.iloc[:100]
xTitanicInstances_5Features_Test_10 = xTitanicInstances_5Features.iloc[-20:]
yTitanicCategories_Test_10 = yTitanicCategories.iloc[-20:]

## scale:
from sklearn.preprocessing import StandardScaler
sc = StandardScaler()
xTrainScaled = pd.DataFrame(sc.fit_transform(xTitanicInstances_5Features_100),index=xTitanicInstances_5Features_100.index)
xTestScaled = pd.DataFrame(sc.transform(xTitanicInstances_5Features_Test_10),index=xTitanicInstances_5Features_Test_10.index)

# yPredicted=predict(xTestInstances, xTrainInstances,yTrainCategories,distanceMetric,k)
yPredicted_scaled_manhattan_3 = predict(xTestScaled,xTrainScaled,yTitanicCategories_100,calc_manhattan_vector_dist,3)
yPredicted_scaled_euclidean_3 = predict(xTestScaled,xTrainScaled,yTitanicCategories_100,calc_euclidean_vector_dist,3)
acc_scaled_manhattan_3 = calc_accuracy(yPredicted_scaled_manhattan_3,yTitanicCategories_Test_10)
acc_scaled_euclidean_3 = calc_accuracy(yPredicted_scaled_euclidean_3,yTitanicCategories_Test_10)
print ('(Scaled) manhattanDist,k=3, acc: %f' %(acc_scaled_manhattan_3))
print ('(Scaled) euclideanDist,k=3, acc: %f' %(acc_scaled_euclidean_3))
assert acc_scaled_manhattan_3 == 0.65,"wrong value for accuracy of (Scaled) knn, with k=3, dist='manhattan'"
assert acc_scaled_euclidean_3 == 0.60,"wrong value for accuracy of (Scaled) knn, with k=3, dist='euclidean'"
# --------------------------------------------------------
# if no exception was thrown by this point, this test was successful:
print ("----> The 'calc_accuracy' test passed successfully :-) \n")

print ('note: Additional Tests might be executed on our side')  

```

Output:
```
testing implementation of the 'calc_accuracy' method ...
(Scaled) manhattanDist,k=3, acc: 0.650000
(Scaled) euclideanDist,k=3, acc: 0.600000
----> The 'calc_accuracy' test passed successfully :-) 

note: Additional Tests might be executed on our side
```

<!-- grade_id: cell-909ea12646d6ce9b, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## addtional tests - predict method implementation - extra test
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_accuracy' method, these tests are not visble to the students ...") 
```

Output:
```
testing implementation of the 'calc_accuracy' method, these tests are not visble to the students ...
```