<H1> Assignment 3 - Naive Bayes</H1><br/>
<table align='left'><tr><td><img src="./images/naiveBayesClassifier.jpeg" alt="Naive Bayes Classifier" style="width: 450px;text-align:left"/></td></tr></table>


<h2>The Iris flowers Dataset</h2>
<h3><u/>Classes</u>- types of Iris flower</h3><br/><br/>
<table align='left'>
    <tr><div style="text-align:left;font-size:125%"><u>There are 3 possible flower types:</u></div></tr>
    <tr><th>class 0 - Iris Setosa</th><th>class 1 - Iris Virginica</th><th>class 2 - Iris Versicolor</th></tr>
    <tr>
        <td><img src="./images/Iris-versicolor.jpg" alt="Iris Versicolor" style="width: 200px;text-align:left"/></td>
        <td><img src="./images/iris_setosa.jpg" alt="Iris Setosa" style="width: 200px;text-align:left"/></td>
        <td><img src="./images/iris_virginica.jpg" alt="Iris Virginica" style="width: 200px;text-align:left"/></td>
    </tr>
</table> 

<h3><u>Feature Set</u></h3>
<table align='left' width="200">
    <tr><td>
        <img src="./images/iris_petal_sepal.png" alt="Iris" width="200" align='left'/>
    </td></tr>
    <tr><td align='left'>
        <div style="text-align:left;font-size:125%"><u>Dataset includes 4 features</u>:</div><br/>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>sepal-length</li>
            <li>sepal-width</li>
            <li>petal-length</li>
            <li>petal-width</li>
        </ul>
    </td></tr>
</table>    

## General instructions:
Please Read the following instructions, and make sure you follow all steps, before submiting your assignment via moodle:
- Put your code one line after the '# YOUR CODE HERE' remark
- You must remove the 'raise NotImplementedError()' exception raising line (which appears a line after the above remark).
    * if you do not remove this line, it will be a sign that you didn't implement that code, and we won't be able to check your work
- Do NOT remove any other line in this notebook    
- you also need to implement the two functions 'myName', 'myId'
    * myName - you need to return your full name as a string
    * myId - you need to return your ID number as a an integer    
- When you want to check your work, select the 'Cell' --> 'Run All' menu
    * before performing your filnal test, we suggest to clear previous output (by selecting 'Cell' --> 'All output' --> 'Clear' menu) and then reperform the final execution ('Cell' --> 'Run All') of the code.
    * after performing 'Run All', make sure there are no exceptions thrown and that the output is as you expected.
- Don't forget to save your work, by clicking the 'save' icon (in the upper left part of the screen), or by selecting the 'File' --> 'Save and checkpoint' menu item. 
- Do NOT change the file name of this notebook
- The work is done individualy, for each student
- Each student needs to submit his/her assignment through moodle.
- Submit the python notebook only (not any additional possible files) 
<br/><br/>
Good Luck :-)<br/>
The courses staff

## Assignmet 3 - Naive Bayes (total of 10 points)
<u>The assignment consists of 6 parts</u>:<br/>
* a preceding step - import packages
* Part 1 - Student information methods ----> needs stundent implementation
* Part 2 - load the Iris dataset ----> run only
* Part 3 - Training a (Guassian) Naive bayes model  ----> needs stundent implementation - 5 point
  * Part 3a - Calculate Priors ----> needs stundent implementation - total 1 point
  * Part 3b - Calculate Gaussian Likelihood 'mean' parameter ---->  needs stundent implementation - total 1 point
  * Part 3c - Calculate Gaussian Likelihood 'std' parameter ---->  needs stundent implementation - total 1 point
  * Part 3d - the "fit" method (main training flow) ----> needs stundent implementation - total 2 points
* Part 4 - predict  ----> needs stundent implementation- total 4 points
  * Part 4a - Calculate Guassian likelihood probability ----> run only
  * Part 4b - Calculate a posteriori probabilities ---->  needs stundent implementation - total 2 points
  * Part 4c - predict (main flow) ---->  needs stundent implementation - total 2 points
* Part 5 - evaluation - calculate error rate  ---->  needs stundent implementation - total 1 point

## Preceding Step - import packages
This step is necessary in order to use external packages. <br/><br/>

```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------
import pandas as pd
import numpy as np
import math
import sys
import os
from random import randrange
from random import seed
```

<h3>Part 1 - Student information methods:</h3><br/>
The following 2 cells consist of 2 methods which aim to validate your details.<br/>
Please <b>make sure to implement <u>both</u> of them</b><br/><br/>
<table align='left'>
    <tr><td>
        <img src="./images/id-card.png" alt="name and id" width="150" align='left'/>
    </td></tr>
</table>  

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myName</b><br/>
<pre>
The following is expected:
--- the method needs to return your full name (as a string)
    For example: assume your name is John Smith, you should write:
    return 'John Smith'
------------
return value:
- your full name (as a string)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myName-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myName
# --------------------------------------------------------
def myName():
    # ADD the return statement in the LINE AFTER: '# YOUR CODE HERE'
    # YOUR CODE HERE
    return "Hen Haim Someh"
```

<u>Part 1 Student information methods</u>:<br/>
<u>method name</u>: <b>myId</b><br/>
<pre>
The following is expected:
--- the method needs to return your ID number (as an integer number)
    For example: assume your ID number is 1234, you should write:
    return 1234
------------
return value:
- your ID number (as an integer number)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the 'YOUR CODE HERE' line, add your code
---------------------
</pre>

<!-- grade_id: myId-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL
# --------------------------------------------------------

# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------

# --------------------------------------------------------
# method name: myId
# --------------------------------------------------------
def myId():
    # YOUR CODE HERE
    return 313468654
```

<u>Part 1 Student information method - test</u><br/>
* Test of 'myName' and 'myId' methods implementation  
* It tests the correctness their implementation

<u>Important Notes</u>:<br/>
* This test cell MUST NOT raise any exception
* It must pass all tests 
* Do NOT change or delete this cell

```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY CELL
# --------------------------------------------------------

# --------------------------------------------------------
# Test:
## the 'myName' and the 'myId' methods implementation 
# --------------------------------------------------------
aName = myName()
aId = myId()
assert aName is not None, 'no return student name'
assert type(aName) is str, "name is not a string: %r" % aName
print ("'myName' method validation is successfull!")
print ("your name is: %s" %(aName))
print ("--------------------------------------")
assert aId is not None, 'no return student id'
assert type(aId) is int, "id is not an integer: %r" % aId
print ("'myId' method validation is successfull!")
print ("your id is: %d" %(aId))
## END OF 'myName' and 'myId' implementation validation
# --------------------------------------------------------
```

Output:
```
'myName' method validation is successfull!
your name is: Hen Haim Someh
--------------------------------------
'myId' method validation is successfull!
your id is: 313468654
```

<h3>Part 2 - Loading dataset - run only</h3>
        <ul style="list-style-type:circle;text-align:left;font-size:115%">
            <li>load the Iris train & test datasets</li>
        </ul>
<img src="./images/load_dataframe.jpg" alt="load dataframe" width="100" align='left'/>


```python
# --------------------------------------------------------
# NO IMPLEMENTAION - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
### The Following section loads the iris dataset directly from sklean
# --------------------------------------------------------
# --- from sklearn import datasets
# --- import os
# --------------------------------------------------------
# --- data_iris = datasets.load_iris()
# --- dataset = pd.DataFrame(data_iris['data'], columns=data_iris['feature_names'])
# --- dataset['target'] = data_iris['target']
# --- dataset_shuffle = dataset.sample(frac=1, replace=False,random_state=41).reset_index()
# --- dataset_shuffle.iloc[:120].to_csv('data'+os.sep+'iris_train_set.csv',index=False)
# --- dataset_shuffle.iloc[120:].to_csv('data'+os.sep+'iris_test_set.csv',index=False)
# --------------------------------------------------------
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------------------------------------------
# display the iris class information:
# --------------------------------------------------------
iris_classes = np.sort(np.unique(y_train.values))
iris_class_names =  ('Setosa','Versicolor','Virginica')
str_classes = ["%d (representing '%s')" %(iris_class,iris_class_names[iris_class]) for iris_class in iris_classes]
print ('There are %d Iris classes' %(len(str_classes)))
print ('iris classes:')
print(str_classes)
# --------------------------------------------------------
# display the feature information:
# --------------------------------------------------------
feature_names = [col for col in x_train.columns]
print ('\nThere are %d features in the feature set (each feature vector has a value for every feature)' %(len(feature_names)))
print('Feature names:')
print(feature_names)
# --------------------------------------------------------
# display the first few rows of the dataset vectors:
# --------------------------------------------------------
print('\nThe train-set includes %d instances (aka feature vectors)' %(x_train.shape[0]))
print('first few instances:')
print(x_train.head())
print('\nfirst few corresponding categories:')
print(y_test.head())
```

Output:
```
There are 3 Iris classes
iris classes:
["0 (representing 'Setosa')", "1 (representing 'Versicolor')", "2 (representing 'Virginica')"]

There are 4 features in the feature set (each feature vector has a value for every feature)
Feature names:
['sepal length (cm)', 'sepal width (cm)', 'petal length (cm)', 'petal width (cm)']

The train-set includes 120 instances (aka feature vectors)
first few instances:
       sepal length (cm)  sepal width (cm)  petal length (cm)  \
index                                                           
119                  6.0               2.2                5.0   
128                  6.4               2.8                5.6   
135                  7.7               3.0                6.1   
91                   6.1               3.0                4.6   
112                  6.8               3.0                5.5   

       petal width (cm)  
index                    
119                 1.5  
128                 2.1  
135                 2.3  
91      
```

## Part 3 - Training a (Guassian) Naive Bayes model

<img src="./images/bayes.PNG" alt="Naive Bayes Classifier" align='left'/>

### Part 3 - Training a (Guassian) Naive bayes model ----> needs stundent implementation - 5  points
  * Part 3a - Calculate Priors ----> needs stundent implementation - total 1 point
  * Part 3b - Calculate Gaussian Likelihood 'mean' parameter ----> needs stundent implementation - total 1 point
  * Part 3c - Calculate Gaussian Likelihood 'std' parameter ----> needs stundent implementation - total 1 point
  * Part 3d - the "fit" method (main training flow) ----> needs stundent implementation - total 2 point

#### Part 3a - Calculate Priors 
for each class (0,1,2) you need to calculate the prior.<br/>
* <b>prior(y=0)</b>=p(y=0)=count(y=0 in train-set)/count(number-of-instances in train-set)
* <b> do this for each class (0,1,2) </b>

<u>Part 3a - Calculate class priors </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_category_priors</b><br/>

<pre>
Change the following function to return a (regular python) list 
containing the 3 classes prirors (one for each category).
------------
input parameter:
- yTrain - the train-set categories
------------
return value:
-  a (regular python) list, containing the priors of the classes
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: calcCategoryPriors-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_category_priors
# --------------------------------------------------------
def calc_category_priors(yTrain):
    unique_classes = np.sort(np.unique(yTrain.values))
    # YOUR CODE HERE
    priorsList = yTrain.value_counts()
    numberOfInstances = len(yTrain)   
    priorsArray = [float(priorsList[0]/numberOfInstances),
                   float(priorsList[1]/numberOfInstances),
                   float(priorsList[2]/numberOfInstances)]
    return priorsArray
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-calcCategoryPriors, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_category_priors' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_category_priors' method ...\n") 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arrPriors = calc_category_priors(y_train)
priorClass_0 = arrPriors[0]
assert int(priorClass_0*100)==34,'wrong prior for class-0'
print ('prior for category 0: %f' %(priorClass_0))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_category_priors' test passed successfully :-) \n")

```

Output:
```
testing implementation of the 'calc_category_priors' method ...

prior for category 0: 0.341667

----> This 'calc_category_priors' test passed successfully :-)
```

#### Part 3b - Calculate Gaussian Likelihood 'mean' parameter:
for each feature calculate mean value for the feature in each for each of the class values (0,1,2) seperatly.<br/><br/>
<b>for class 0 (y=0)</b> take the rows consisting 'target' value of 0, and calculate the mean <br/>
To calculate mean use: dataframe[colName].mean() <br/><br/>
<b>Do this for each feature for each class</b>

<u>Part 3b - Calculate the 'mean' element out of likelihhod probability</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_mean_likelihood</b><br/>

<pre>
Change the following function to return mean of the values of each feature
this should be done seperatly for each class value (in the training set)
    * create a meanDataframe, including the features as colums and classes as indexes
    * the cells consists of the mean of distribution of feature values (given a specific class)
------------
input parameter:
- training_set - the train-set dataframe
------------
return value:
-  a dataframe, consisting of the feature names as colums and classes as indexes
   the cells consists of the mean of distribution of feature values (given a specific class)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: calcMeanLikelihood-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_mean_likelihood
# --------------------------------------------------------
def calc_mean_likelihood(training_set):
    # YOUR CODE HERE
    dataFrameMeanArray = pd.DataFrame({ '0' : training_set[training_set['target'] == 0].mean(),
                                        '1' : training_set[training_set['target'] == 1].mean(),
                                        '2' : training_set[training_set['target'] == 2].mean()})
    dataFrameMeanArray = dataFrameMeanArray.T
    return dataFrameMeanArray
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-calcMeanLikelihood, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_mean_likelihood' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_mean_likelihood' method ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
meanLiklihoodDf = calc_mean_likelihood(train_set)
likelihood_petalLength_class1 = meanLiklihoodDf.iloc[1,2]
assert int(likelihood_petalLength_class1*100)==427,'wrong mean likelihood estimation for petalLength_class1'
print ('likelihood for the mean of petal length for category 1 is estimated as: %f' %(likelihood_petalLength_class1))
# --------------------- 
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_mean_likelihood' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'calc_mean_likelihood' method ...

likelihood for the mean of petal length for category 1 is estimated as: 4.279487

----> This 'calc_mean_likelihood' test passed successfully :-)
```

#### Part 3c - Calculate Gaussian Likelihood 'std' parameter:
for each feature calculate std value for the feature in each for each of the class values (0,1,2) seperatly.<br/><br/>
<b>for class 0 (y=0)</b> take the rows consisting 'target' value of 0, and calculate the std <br/>
To calculate std use: dataframe[colName].std() <br/><br/>
<b>Do this for each feature for each class</b>

<u>Part 3c - Calculate the 'std' element out of likelihhod probability</u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_std_likelihood</b><br/>

<pre>
Change the following function to return std (standard diviation) of the values of each feature
this should be done seperatly for each class value (in the training set)
    * create a stdDataframe, including the features as colums and classes as indexes
    * the cells consists of the std of distribution of feature values (given a specific class)
------------
input parameter:
- training_set - the train-set dataframe
------------
return value:
-  a dataframe, consisting of the feature names as colums and classes as indexes
   the cells consists of the std of distribution of feature values (given a specific class)
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: calcStdLikelihood-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_std_likelihood
# --------------------------------------------------------
def calc_std_likelihood(training_set):
    # YOUR CODE HERE
    dataFrameStdArray = pd.DataFrame({ '0' : training_set[training_set['target'] == 0].std(),
                                       '1' : training_set[training_set['target'] == 1].std(),
                                       '2' : training_set[training_set['target'] == 2].std()})
    dataFrameStdArray = dataFrameStdArray.T
    return dataFrameStdArray
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-calcStdLikelihood, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_std_likelihood' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_std_likelihood' method ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
stdLiklihoodDf = calc_std_likelihood(train_set)
likelihood_petalLength_class1 = stdLiklihoodDf.iloc[1,2]
assert int(likelihood_petalLength_class1*1000)==501,'wrong std likelihood estimation for petalLength_class1'
print ('likelihood for the std of petal length for category 1 is estimated as: %f' %(likelihood_petalLength_class1))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_std_likelihood' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'calc_std_likelihood' method ...

likelihood for the std of petal length for category 1 is estimated as: 0.501146

----> This 'calc_std_likelihood' test passed successfully :-)
```

#### Part 3d - the fit method:
The fit method uses the previous 3 methods for a full (Gaussian) Naive Bayes model training step.<br/>


<u>Part 3d - <b>fit</b> - Naive Bayes main training flow</u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>fit</b><br/>

<pre>
Change the following function to return the 3 products of the Naive Bayes training.
Note: this step is called the 'fit' api in machine learning.
The method needs to calculate and return the products from the 3 previous methods you've 
implemented in the previous cells, including calc_category_priors, calc_mean_likelihood, calc_std_likelihood.
Please check the expected output for these methods.
------------
input parameter:
- training_set - the train-set dataframe
------------
return values (comma seperated):
- lst_class_prior    - output from calc_category_priors
- df_mean_likelihood - output from calc_mean_likelihood
- df_std_likelihood  - output from calc_std_likelihood
return statememt (assuming the same variable names):
return lst_class_prior,df_mean_likelihood,df_std_likelihood
-----
note the return statement should be done as following (assuming these are the names of the variables):
return lst_class_prior,df_mean_likelihood,df_std_likelihood
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)



# ------------------------------------------
```

<!-- grade_id: fit-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: fit
# --------------------------------------------------------
def fit(training_set):
    # YOUR CODE HERE
    return (calc_category_priors(training_set['target']), calc_mean_likelihood(training_set), calc_std_likelihood(training_set))
```

<!-- grade_id: test-fit, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'fit' method #1 ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arrPriors, meanLiklihoodDf, stdLiklihoodDf = fit(train_set)
priorClass_2 = arrPriors[2]
likelihood_mean_sepalWidth_class2 = meanLiklihoodDf.iloc[2,1]
likelihood_std_sepalWidth_class2 = stdLiklihoodDf.iloc[2,1]
assert int(priorClass_2*100)==33,'wrong prior for class-2'
print ('prior for category 2: %f' %(priorClass_2))
assert int(likelihood_mean_sepalWidth_class2*10)==29,'wrong mean likelihood estimation for sepalWidth_class2'
print ('likelihood for the mean of sepal width for category 2 is estimated as: %f' %(likelihood_mean_sepalWidth_class2))
assert int(likelihood_std_sepalWidth_class2*100)==31,'wrong std likelihood estimation for sepalWidth_class2'
print ('likelihood for the std of sepal width for category 2 is estimated as: %f' %(likelihood_std_sepalWidth_class2))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_std_likelihood' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'fit' method #1 ...

prior for category 2: 0.333333
likelihood for the mean of sepal width for category 2 is estimated as: 2.945000
likelihood for the std of sepal width for category 2 is estimated as: 0.314561

----> This 'calc_std_likelihood' test passed successfully :-)
```

<!-- grade_id: test-fit-2, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'fit' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'fit' method #2, this test is not visble to the students ...\n") 
# --------------------- 
```

Output:
```
testing implementation of the 'fit' method #2, this test is not visble to the students ...
```

### Part 4 - predict ----> needs stundent implementation- total 4 points
  * Part 4a - Calculate Guassian likelihood probability ----> run only
  * Part 4b - Calculate a posteriori probabilities ----> needs stundent implementation - total 2 points
  * Part 4c - prdict ----> needs stundent implementation - total 2 point

<img src="./images/bayes.PNG" alt="Naive Bayes Classifier" align='left'/>

###  Part 4a  - the 'calcGaussianProb' method:
The 'calcGaussianProb' method uses the training methods and returns the Gaussian probablilty <br/>
of that feature value (for a specifc class).<br/>
##### Note: you NEED to run this cell (but NOT implement anything)

<img src="./images/gausianProb.PNG" alt="Gausian likelihood probability" align='left'/>

```python
# --------------------------------------------------------
# YOU MUST RUN THIS NOTEBOOK CELL
# --------------------------------------------------------
"""
given a specific feature value and the trained mean & std (per a specific class) 
We assume normal (Guassian distribution) and we return the density value 
or the Gaussian Probability for that given value
Note: the input parameters are all numbers (scalars)
"""
def calc_gaussian_prob(x_feature_val, feature_mean, feature_std):
    exponent = np.exp(-((x_feature_val-feature_mean)**2 / (2 * feature_std**2 )))
    return (1 / ((2 * np.pi)**(1/2) * feature_std)) * exponent
```

### Part 4b - the 'calc_aposterior_probs' method:
The 'calc_aposterior_probs' method uses the training parameters to predict the a posteriori probability <br/>
for every test instance, per class <br/>
##### Note: you NEED to run this cell (but NOT implement anything)

<u>Part 4b - calculate Aposterior probabilities</u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>calc_aposterior_probs</b><br/>

<pre>
The following function should create a probability matrix to store the results.
To do that create a dataframe with the categories as columns and test indexes ans indexes.
The value of the cells should be a multiplication between the class prior and the Gausian probabilities of all the features (using the  'calcGaussianProb' method) for the specific test feature vectore (i.e., Gausian probabilities for feature values of every row).
The outcome is the Aposterior probabilities for each feature vector (one for each class). 
Please check the expected output for these methods.
------------
input parameter:
- x_test - the test-set instances
- arr_trained_class_priors - the trained priors of the classes
- df_trained_mean - the dataframe of the trained means for each feature per each class
- df_trained_std - the dataframe of the trained stds for each feature per each class
- categories - the possible categories simple python list.
               remark: in the iris case includes [0,1,2]
------------
return value:
------------
- df_aposterior_probs -  a dataframe containing posterior probabilities per class, per test feature vector.
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)


# ------------------------------------------
```

<!-- grade_id: calc_aposterior_probs-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_aposterior_probs
# --------------------------------------------------------
def calc_aposterior_probs(x_test, arr_trained_class_priors, df_trained_mean, df_trained_std, categories):
    num_classes = len(categories)
    df_prob_per_test_inst_per_class = pd.DataFrame(np.zeros((x_test.shape[0], num_classes)), columns=categories, index=x_test.index)
    # YOUR CODE HERE
    
    for category in categories:
        df_prob_per_test_inst_per_class[category] = arr_trained_class_priors[category]
        
        for row in range(0, x_test.shape[0]):    
            for column in range(0, x_test.shape[1]):
                df_prob_per_test_inst_per_class.iloc[row, category] *= calc_gaussian_prob(x_test.iloc[row, column],
                                                                                          df_trained_mean.iloc[category, column],
                                                                                          df_trained_std.iloc[category, column])        
    return df_prob_per_test_inst_per_class
    
```

<!-- grade_id: test-calc_aposterior_probs, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_aposterior_probs' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_aposterior_probs' method #1 ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arr_trained_priors = calc_category_priors(y_train)
trained_mean_df = calc_mean_likelihood(train_set)
trained_std_df = calc_std_likelihood(train_set)
unique_categories = np.sort(np.unique(y_train.values))

df_prob_per_test_inst_per_class = calc_aposterior_probs(x_test, arr_trained_priors, trained_mean_df, trained_std_df, unique_categories)
assert isinstance(df_prob_per_test_inst_per_class,pd.DataFrame),'calc_aposterior_probs should return a dataframe'
assert df_prob_per_test_inst_per_class.shape[0]==x_test.shape[0],'calc_aposterior_probs should return a dataframe with %d rows' %(x_test.shape[0])
assert df_prob_per_test_inst_per_class.shape[1]==len(unique_categories),'calc_aposterior_probs should return a dataframe with %d columns' %(len(unique_categories))
assert int( df_prob_per_test_inst_per_class.loc[103,2]*1000)==207,'wrong aposterior probability for class: 2, index: 103'
print('aposterior probability of row with index 103 (2nd row), class 2: %f' %(df_prob_per_test_inst_per_class.loc[103,2]))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_aposterior_probs' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'calc_aposterior_probs' method #1 ...

aposterior probability of row with index 103 (2nd row), class 2: 0.207558

----> This 'calc_aposterior_probs' test passed successfully :-)
```

<!-- grade_id: test-calc_aposterior_probs-2, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_aposterior_probs' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_aposterior_probs' method #1 ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arr_trained_priors = calc_category_priors(y_train)
trained_mean_df = calc_mean_likelihood(train_set)
trained_std_df = calc_std_likelihood(train_set)
unique_categories = np.sort(np.unique(y_train.values))

df_prob_per_test_inst_per_class = calc_aposterior_probs(x_test, arr_trained_priors, trained_mean_df, trained_std_df, unique_categories)
#print (df_prob_per_test_inst_per_class.where(df_prob_per_test_inst_per_class>0.01).head())
assert isinstance(df_prob_per_test_inst_per_class,pd.DataFrame),'calc_aposterior_probs should return a dataframe'
assert df_prob_per_test_inst_per_class.shape[0]==x_test.shape[0],'calc_aposterior_probs should return a dataframe with %d rows' %(x_test.shape[0])
assert df_prob_per_test_inst_per_class.shape[1]==len(unique_categories),'calc_aposterior_probs should return a dataframe with %d columns' %(len(unique_categories))
assert int( df_prob_per_test_inst_per_class.loc[44,0]*1000)==13,'wrong aposterior probability for class: 0, index: 44'
print('aposterior probability of row with index 44 (1st row), class 0: %f' %(df_prob_per_test_inst_per_class.loc[44,0]))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_aposterior_probs' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'calc_aposterior_probs' method #1 ...

aposterior probability of row with index 44 (1st row), class 0: 0.013257

----> This 'calc_aposterior_probs' test passed successfully :-)
```

### Part 4c- the 'predict' method:
The 'predict' method is a Guasian Naive Bayes classifier <br/>
Use the previous method (calc_aposterior_probs), to calculate  a posteriori probabilites <br/>
for every test instance.
And calculate the most probable class for each test instance.

<u>Part 4c - predict </u> ----> needs stundent implementation - 2 point<br/>
<u>method name</u>: <b>predict</b><br/>

<pre>
Change the following function to return the predicted categories of the test set examples
one category for each correcponding test instance
this should be done seperatly for each class value (in the training set), using the Gaussian Naive Bayes.
Note: For each test veture vector, use the Aposterior Probability for each class. 
      The predicted class (in the example level), should be the one with the maximum probability.
------------
input parameter:
- x_test - the test-set instances
- arr_trained_class_priors - the trained priors of the classes
- df_trained_mean - the dataframe of the trained means for each feature per each class
- df_trained_std - the dataframe of the trained stds for each feature per each class
- categories - the possible categories simple python list.
               remark: in the iris case includes [0,1,2]
------------
return value:
------------
- y_predicted -  a series of containing all predicted class values per test instance
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)


# ------------------------------------------
```

<!-- grade_id: predict-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 2
# --------------------------------------------------------

# --------------------------------------------------------
# method name: predict
# --------------------------------------------------------
def predict(x_test, arr_trained_class_priors, df_trained_mean, df_trained_std, categories):
    # YOUR CODE HERE
    dataFramePosterior = calc_aposterior_probs(x_test, arr_trained_class_priors, df_trained_mean, df_trained_std, categories)
    return dataFramePosterior.idxmax(axis = 1)
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-predict, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict' method #1 ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arr_trained_priors = calc_category_priors(y_train)
trained_mean_df = calc_mean_likelihood(train_set)
trained_std_df = calc_std_likelihood(train_set)
unique_categories = np.sort(np.unique(y_train.values))
y_predicted = predict(x_test, arr_trained_priors, trained_mean_df, trained_std_df, unique_categories)

assert isinstance(y_predicted,pd.Series),'predict should return a series'
assert y_predicted.shape[0]==y_test.shape[0],'predict should return a series with %d rows' %(y_test.shape[0])
assert y_predicted.loc[44]==0,'wrong predicted value for index 44 (first value), expects: 0'
assert y_predicted.iloc[4]==1,'wrong predicted value for 5th value, expects: 1'
print ('y_predicted first few rows:')
print (y_predicted.head())

# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'predict' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'predict' method #1 ...

y_predicted first few rows:
index
44     0
103    2
5      0
20     0
69     1
dtype: int64

----> This 'predict' test passed successfully :-)
```

<!-- grade_id: test-predict-2, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'predict' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'predict' method #2, this test is not visble to the students ...\n") 
# --------------------- 
```

Output:
```
testing implementation of the 'predict' method #2, this test is not visble to the students ...
```

### Part 5 - evaluation (error rate) ----> needs stundent implementation - 1 point
<img src="./images/inaccuracy.jpg" alt="Accuracy" width="300" align='left'/>

<u>Part 5 - Evaluation (error rate) </u> ----> needs stundent implementation - 1 point<br/>
<u>method name</u>: <b>calc_error_rate</b><br/>

<pre>
The following is expected:
- the method needs to return the error rate (incorrectly predicted ratio)
------------
input parameter:
- y_predicted -  a series of containing all predicted class values per test instance
- y_test - a series of containing all actual class values per test instance
------------
return value:
------------
- errorRateVal - the error rate value
---------------------
---------------------
--- General notes:
Don't forget to remove the 'raise NotImplementedError()' line.
In the line following the '# YOUR CODE HERE' line, add your code
---------------------
</pre>

```python
# ------------------------------------------
# Add assistance code here IF NEEDED:
#  - RUN AFTER IMPLEMENTATION (if used)


# ------------------------------------------
```

<!-- grade_id: calc_error_rate-answer, points: 0, grade: false, solution: true -->
```python
# --------------------------------------------------------
# IMPLEMENTATION CELL - RUN AFTER IMPLEMENTATION
# --------------------------------------------------------

# --------------------------------------------------------
# Grading - Maximum points: 1
# --------------------------------------------------------

# --------------------------------------------------------
# method name: calc_error_rate
# --------------------------------------------------------
def calc_error_rate(y_predicted,y_test):
    # YOUR CODE HERE
    errorRate = 0   
    for i in range (0, y_predicted.shape[0]):
        if (y_test.iloc[i] != y_predicted.iloc[i]):
            errorRate = errorRate + 1
            
    errorRate = errorRate / y_predicted.shape[0]
    return errorRate
```

```python
# ---------------- internal-students-tests ---------------------------
# - add your tests here IF NEEDED:
# Note: any additional tests are for your use only.
# Note: We depend only on tests we wrote(and disregard any tests you add).
# ------


# --------------------------------------------------------
```

<!-- grade_id: test-error_rate, points: 1, grade: true, solution: false -->
```python
# --------------------------------------------------------
# TEST ONLY CELL - DO NO CHANGE - RUN ONLY
# --------------------------------------------------------

# --------------------------------------------------------
# Test the:
## 'calc_error_rate' method implementation 
# --------------------------------------------------------

#--------------------------------------------------------
print ("testing implementation of the 'calc_error_rate' method ...\n")
# --------------------- 
# --------------------- 
train_set = pd.read_csv('data'+os.sep+'iris_train_set.csv').set_index('index')
x_train = train_set.drop('target', axis=1)
y_train = train_set['target']
# ----
test_set  = pd.read_csv('data'+os.sep+'iris_test_set.csv').set_index('index')
x_test = test_set.drop('target', axis=1)
y_test = test_set['target']
# --------------------- 
arr_trained_priors = calc_category_priors(y_train)
trained_mean_df = calc_mean_likelihood(train_set)
trained_std_df = calc_std_likelihood(train_set)
unique_categories = np.sort(np.unique(y_train.values))
y_predicted = predict(x_test, arr_trained_priors, trained_mean_df, trained_std_df, unique_categories)
error_rate = calc_error_rate(y_predicted,y_test)
assert float(int(error_rate*1000))/1000 == 0, "error rate should be 0"
print("Error Rate: {}".format(error_rate))
# --------------------------------------------------------
# if no exception was thrown by this point, the student side tests were successful:
print ("\n----> This 'calc_error_rate' test passed successfully :-) \n")
```

Output:
```
testing implementation of the 'calc_error_rate' method ...

Error Rate: 0.0

----> This 'calc_error_rate' test passed successfully :-)
```